ANKRD5: a key component of the axoneme required for sperm motility and male fertility

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public review):

Summary:

Asthenospermia, characterized by reduced sperm motility, is one of the major causes of male infertility. The "9 + 2" arranged MTs and over 200 associated proteins constitute the axoneme, the molecular machine for flagellar and ciliary motility. Understanding the physiological functions of axonemal proteins, particularly their links to male infertility, could help uncover the genetic causes of asthenospermia and improve its clinical diagnosis and management. In this study, the authors generated Ankrd5 null mice and found that ANKRD5-/- males exhibited reduced sperm motility and infertility. Using FLAG-tagged ANKRD5 mice, mass spectrometry, and immunoprecipitation (IP) analyses, they confirmed that ANKRD5 is localized within the N-DRC, a critical protein complex for normal flagellar motility. However, transmission electron microscopy (TEM) and cryo-electron tomography (cryo-ET) of sperm from Ankrd5 null mice did not reveal any structural abnormalities.

Strengths:

The phenotypes observed in ANKRD5-/- mice, including reduced sperm motility and male infertility, are conversing. The authors demonstrated that ANKRD5 is an N-DRC protein that interacts with TCTE1 and DRC4. Most of the experiments are thoughtfully designed and well executed.

Weaknesses:

The cryo-FIB and cryo-ET analyses require further investigation, as detailed below. The molecular mechanism by which the loss of ANKRD5 affects sperm flagellar motility remains unclear. The current conclusion that Ankrd5 knockout reduces axoneme stability is not well-supported. Specifically, are other axonemal proteins diminished in Ankrd5 knockout sperm? Conducting immunofluorescence analyses and revisiting the quantitative proteomics data may help address these questions.

Reviewer #2 (Public review):

Summary:

The manuscript investigates the role of ANKRD5 (ANKEF1) as a component of the N-DRC complex in sperm motility and male fertility. Using Ankrd5 knockout mice, the study demonstrates that ANKRD5 is essential for sperm motility and identifies its interaction with N-DRC components through IP-mass spectrometry and cryo-ET. The results provide insights into ANKRD5's function, highlighting its potential involvement in axoneme stability and sperm energy metabolism.

Strengths:

The authors employ a wide range of techniques, including gene knockout models, proteomics, cryo-ET, and immunoprecipitation, to explore ANKRD5's role in sperm biology.

Weaknesses:

(1) Limited Citations in Introduction: Key references on the role of N-DRC components (e.g., DRC1, DRC2, DRC3, DRC5) in male infertility are missing, which weakens the contextual background.

(2) Lack of Functional Insights: While interacting proteins outside the N-DRC complex were identified, their potential roles and interactions with ANKRD5 are not adequately explored or discussed.

(3) Mitochondrial Function Uncertainty: Immunofluorescence suggests possible mitochondrial localization for ANKRD5, but experiments on its role in energy metabolism (e.g., ATP production, ROS) are insufficient, especially given the observed sperm motility defects.

(4) Glycolysis Pathway Impact: Proteomic analysis indicates glycolysis pathway disruptions in Ankrd5-deficient sperm, but the link between these changes and impaired motility is not well explained.

(5) Cryo-ET Data Limitations: The structural analysis of the DMT lacks clarity on how ANKRD5 influences N-DRC or RS3. The low quality of RS3 data hinders the interpretation of ANKRD5's impact on axoneme structure.

(6) Discussion of Findings: The manuscript could benefit from a deeper discussion on the broader implications of ANKRD5's interactions and its role in sperm energy metabolism and motility mechanisms.

Reviewer #1 (Recommendations for the authors):

EMD-35210/35211 are 16-nm maps while the Ankrd5 null map is 8-nm repeat. To generate a difference map, the authors should use maps of the same periodicity.

Thank you for your suggestion. We have replaced the old 16-nm maps with an 8nm map and updated the images (Fig. 7). The 8nm repeats DMT density map we used was obtained by summing two 16nm repeats DMTs that were staggered 8nm apart from each other (EMD-35229). The replacement of the 16nm repeats DMT density map with the 8nm repeats DMT density map has no effect on our scientific findings and experimental conclusions.

"We were able to detect the N-DRC structure in WT sperm, but we failed to find the density of N-DRC adjacent to RS3 in Ankrd5 null sperm". Do the authors imply that the N-DRC is lost in Ankrd5 null sperm? To draw a conclusion, they need to compare the 96-nm map of WT sperm axoneme with that of Ankrd5 null sperm axoneme. Quantitative proteomics shows that the levels of most N-DRC components in Ankrd5 null sperm are comparable with those of WT sperm. Why are the quantitative proteomics results not consistent with the structural observation?

We are very sorry for this improper description. Our original description was not rigorous, which led to misunderstanding. Our original intention is to say that the quality of the density map causes the N-DRC to be difficult to recognize, rather than that the N-DRC has disappeared. In addition, attempts to classify 96nm repeats DMT structure during our data processing failed. In the process of classification, we found that the density of RS was not good. So we changed the picture and the description.

We have changed the description in the text: "During the STA process, many particles were misaligned or deformed in the classification results, revealing various degrees of deformation—particularly affecting the B-tube (Figure 9,Fig. S9E). We could retain only ~10% of the DMT particles to obtain the final density map for ANKRD5-KO sperm (Fig. S9E), whereas ~70% were usable in WT dataset as reported previously [59]. The mutant DMT density map also displayed roughness at its periphery, indicating substantial structural heterogeneity (Fig. S9E). Even after discarding a large fraction of deformed particles, the final density map still showed evident artifacts, implying that although the mutant DMT preserves the fundamental features of both tubes, its shape is highly heterogeneous (Fig. S9E). Furthermore, attempts to classify the 96-nm repeats did not yield a clear density for radial spokes (RSs) (Fig. S9F), indicating that ANKRD5 deficiency may affect the stability of other accessory structures, such as RSs [24-26]. In the raw tomograms, RSs in ANKRD5-KO sperm appeared less regularly arranged than those in WT(Fig. S9A and C)."

Figure S9. The states of DMT particles in sperm of Ankrd5-KO mouse. (A) and (C) Tomogram slices of WT and Ankrd5-KO in Dynamo (The data for WT mouse sperm was EMPIARC-200007). DMT and RS are marked with white dashed lines and white arrows, respectively. (B) and (D) Comparison of DMT particle states between WT and Ankrd5-KO in Dynamo. The visual angles of the DMT particles shown in (B) and (D) show that the DMT fibers within the white box in (A) and (B) are divided equally into 10 slices along the direction of the white arrow, respectively. The DMT particle shapes of WT and Ankrd5-KO are marked by white dashed lines on the right of (B) and (D). The white arrow in (D) identifies the junction of A-tube and B-tube that is suspected to be disconnected. (E) Deformed particles discarded in 3D classification and final aligned DMT artifacts. (F) 3D classification of attempted RS locations.

In the process of obtaining DMT with a period of 8nm, we discarded about 90% of the particles (some were mis-aligned particles and some were deformed particles). Although the final DMT density showed complete A-tube and B-tube, both the particles in our calculation process and the projection of the final structure showed strong particle heterogeneity.

Our results show that in ANKRD5-KO mice, the structure of sperm DMT itself has no apparent effect in tube A and tube B, and we found that DMT in the original tomography were not smooth. We speculate that loss of ANKRD5 may reduce the interaction between N-DRC and neighboring DMTs, resulting in nonuniform force on the axoneme during sperm swimming, which may limit our ability to obtain an average structure of the more dynamic components (RS, N-DRC, ODA, IDA). Therefore, when trying to classify 96nm repeat DMTS, we can only see the density of suspected RS3 and RS2, but it is difficult to obtain the confident 96nm repeat DMT density. It is difficult to further discuss the effects of ANKRD5 on RS3 and N-DRC. To test this conjecture, we further classified the density of suspected RS3, and the results obtained exhibited a variety of mixed states (Fig. S9). To avoid confusion, we have already removed the discussion of RS3 and the related images from the original text.

It's not clear whether N-DRC proteins and ODA, IDA, RS proteins are affected in DMT of Ankrd5 null sperm. Immunofluorescence staining would help to resolve this problem.

Thank you for your suggestion. The levels of N-DRC proteins and ODA, IDA, RS were detected by immunofluorescence, and no difference was found between ANKRD5-null sperm and control. We added figure S6 as a new figure and added the following description in red font on page 7 of the article:

Figure S6. Immunofluorescence results of ANKRD5-null sperm and control. DRC11 serves as a marker protein for N-DRC (nexin-dynein regulatory complex), NME5 as a marker for RS (radial spoke), DNALI1 as a marker for IDA (inner dynein arm), and DNAI1 as a marker for ODA (outer dynein arm).

In addition, ODA and RS were also marked in the figure when we further analyzed the Cryo-ET data (Figure 7 and Figure S9).

Does Ankrd5 express in other cilia cells except for sperm?

We stained mouse respiratory cilia using immunofluorescence and found that the protein was also expressed in mouse respiratory cilia. To support this finding, we added Figure S3 as a new figure and included a description in red font on page 6 of the article.

Page 7, "However, in the process of manual selection of DMT fibers, we found that they were not as smooth as WT particles." This description is too subjective. Please show the data.

Thank you for your suggestion. We have added a supplementary figure showing the difference between mutant samples and WT samples during particle picking (Fig. S9).

Abstract, "These findings establish that ANKRD5 is critical for maintaining axoneme stability, "Page 7, "This suggests that the knockout of Ankrd5 may affect the structural stability of the axoneme," I do not see direct evidence that Ankrd5 KO reduces the axoneme stability.

Our phrasing was not sufficiently precise. These findings suggest that ANKRD5 plays a crucial role in limiting the relative sliding between adjacent microtubule doublets during axoneme bending, rather than directly contributing to the stability of the axoneme. This sentence has already been modified in the abstract and marked in red. We have added the description in the text: "These findings suggest that ANKRD5 may weaken the N-DRC’s "car bumper" role, reducing the buffering effect between adjacent DMTs and thereby destabilizing axoneme structures during intense axoneme motility." and "To further investigate the RS, IDA, and ODA structures of the axonemes, we conducted immunofluorescence assays in both Ankrd5-/- mice and the control group. No significant differences were detected between the two groups (Fig. S6)."

Page 8, "but our study offers new perspectives for male contraceptive research". Could the authors expand this a bit - how this study may offer new perspectives for male contraceptive research?

We sincerely appreciate the reviewer's insightful feedback regarding the translational potential of our findings. This is indeed a critical aspect that we sought to highlight. In response, we have added a paragraph on page 9 (marked in red) to further emphasize this point. We have added the description in the text: "The potential for male contraceptive development arises from ANKRD5's critical structural role mediated through its ANK domain, which facilitates interaction with the N-DRC complex in sperm flagella. Recent structural evidence suggests the protein's positively charged surface may engage with glutamylated tubulin in adjacent microtubules[41], presenting a druggable interface. Targeted disruption of this interaction through small-molecule inhibitors could transiently impair sperm motility. Sperm function relies more on ANKRD5 than respiratory cilia, so inhibiting ANKRD5 has less impact on the latter. This makes ANKRD5 a promising drug target. This tissue-specific phenotypic uncoupling is not uncommon among axonemal-associated proteins, such as DNAH17 and IQUB[65，66]."

Abstract, "reveals its interaction with TCTE1 and DRC4/GAS8", please provide the alias symbol DRC5 for TCTE1 for clarity.

Thank you for your suggestion, I have revised the abstract by replacing "TCTE1" with "DRC5/TCTE1" to clarify the alias. The changes have been highlighted in red in the manuscript for easy reference.

Introduction, "Fertilization relies on successful spermatogenesis and normal sperm motility (4), which occurs in the testes." Does spermatogenesis or normal sperm motility occur in the testes?

Thank you for pointing out the ambiguity in the sentence. We have revised the sentence in the Introduction and highlighted it in red as follows: Fertilization relies on successful spermatogenesis and normal sperm motility..

Introduction, "The axoneme exhibits a 9+2 microtubule doublet structure". The description is not accurate. The "2" are singlet microtubules.

Thank you for your suggestion. We have revised the sentence to accurately describe the axoneme structure and highlight in red as follows: The axoneme features a 9+2 architecture, comprising nine doublet microtubules encircling a central pair of singlet microtubules, with the N-DRC forming cross-bridges between adjacent doublets.

Page 4, "control sperm successfully fertilized both cumulus-intact eggs". "control" should be a capital "C".

We thank the reviewer for noting this oversight. The correction has been implemented on page 5 with the term highlighted in red (now reading: "Control sperm successfully fertilized both cumulus-intact eggs"), and we have verified capitalization consistency throughout the manuscript.

Page 6, "applied RELION, M, and other software". "other software" is not an appropriate description, please be precise.

We have revised the description as suggested. Specifically, on page 7, the phrase "and other software" has been replaced with "Dynamo and Warp/M," and this change is highlighted in red for clarity.

Reviewer #2 (Recommendations for the authors):

Several components of the N-DRC complex (e.g., DRC1, DRC2, DRC3, DRC5) have been reported to be associated with male infertility in both humans and mice. However, the introduction lacks proper citations for these studies. Adding these references would provide a more comprehensive background for readers.

Thank you for your suggestion to strengthen the comprehensiveness of the research background by incorporating additional literatures. More literatures related to DRC1, DRC2, DRC3, and DRC5 were cited in the background of this paper. We have rewritten and reorganized the language of the last paragraph of the introduction, and the entire paragraph is highlighted in red. The content of the paragraph is as follows:

"It was previously believed that N-DRC comprised 11 protein components[13，18]. However, a new component CCDC153 (DRC12) was found to interact with DRC1[19]. In situ cryoelectron tomography (cryo-ET) has significantly advanced understanding of the N-DRC architecture in Chlamydomonas, demonstrating that DRC1, DRC2/CCDC65, and DRC4/GAS8 constitute its core framework[16], while proteins DRC3/5/6/7/8/11 associate with this framework and engage with other axonemal complexes[20]. Biochemical experiments corroborate these findings and validate this structural model[12，21，22]. The N-DRC functions between the DMTs to convert sliding into axonemal bending motion by restricting the relative sliding of outer microtubule doublets[23，24，25]. Mutations of N-DRC subunits demonstrate that the structural integrity of the N-DRC is crucial for flagellar movements. Mutations in DRC1, DRC2/CCDC65, and DRC4/GAS8 are linked to ciliary motility disorders, causing primary ciliary dyskinesia (PCD)[12，26]. Biallelic truncating mutations in DRC1 induce multiple morphological abnormalities of sperm flagella (MMAF), including outer DMT disassembly, mitochondrial sheath disorganization, and incomplete axonemal structures in human sperm[22，27，28]. Similarly, CCDC65 loss disrupts N-DRC stability, leading to disorganized axonemes, global microtubule dissociation, and complete asthenozoospermia[12，29]. Homozygous frameshift mutations in DRC3 impair N-DRC assembly and intraflagellar transport (IFT), resulting in severe motility defects despite normal sperm morphology[30，31]. TCTE1 knockout mice maintain normal sperm axoneme structure but show impaired glycolysis, leading to reduced ATP levels, lower sperm motility, and male infertility[32]. Both Drc7 and Iqcg (Drc9) knockout mice exhibit disrupted '9+2' axonemal architecture, sperm immotility, and male infertility[21，33]. Drc7 knockout sperm also display head deformities and shortened tails[21]. While N-DRC is critical for sperm motility, but the existence of additional regulators that coordinate its function remains unclear. Our findings indicate that ANKRD5 (Ankyrin repeat domain 5; also known as ANK5 or ANKEF1) interacts with N-DRC structure, serving as an auxiliary element to facilitate collaboration among DRC members. The absence of ANKRD5 results in diminished sperm motility and consequent male infertility."

While many N-DRC components were identified as interacting with ANKRD5, other proteins outside the N-DRC complex were also detected. Notably, GAS8 (DRC4) ranked 165th among the identified proteins. What are the functions of the higher-ranking proteins, and why do they interact with ANKRD5? Discussing their potential roles would enhance the mechanistic understanding of ANKRD5's function.

We thank the reviewer for highlighting the importance of non-N-DRC proteins interacting with ANKRD5 (ANKEF1). Below, we provide a detailed analysis of the roles and interaction mechanisms of the top-ranked non-N-DRC proteins (Krt77, Rab2a, Gm7429) to elucidate their functional relevance to ANKRD5. We have added the following text to page 6 to clarify and highlight this in red:

As for other proteins in the LC-MS results, KRT77 is a classic protein that maintains cytoskeletal stability. It may enhance the physical connection between the N-DRC and adjacent DMTs through interaction with ANKRD5. Recent studies indicate that ANKRD5, a newly identified component in the distal lobe of the N-DRC, has a positively charged surface, which may facilitate binding to glutamylated tubulin on adjacent DMTs[41]. Thus, KRT77 may also regulate its interaction with ANKRD5 via post-translational modifications (PTMs, e.g., phosphorylation), thereby strengthening sperm resistance to shear forces during flagellar movement. Rab family proteins participate in intraflagellar transport and membrane dynamics. RAB2A may promote targeted transport of ANKRD5 or other N-DRC components to axonemal assembly sites by recruiting vesicles, and its GTPase activity might link cellular signals to ANKRD5-mediated axoneme remodeling. However, the observed signals could be false positives due to nonspecific factors such as electrostatic adsorption, high-abundance protein interference, detergent-induced membrane disruption, or protein aggregation tendencies.

The immunofluorescence localization of ANKRD5-Flag appears more aligned with the mitochondrial sheath rather than the axoneme. There is a finer red fluorescent signal extending from the mitochondrial sheath that might correspond to the axoneme. Could this suggest that ANKRD5 has a functional role in the mitochondria? While the authors measured ROS levels, this might not fully clarify whether ANKRD5 is involved in sperm energy metabolism. Considering the motility defects in Ankrd5 knockout mice, further experiments to explore ANKRD5's potential involvement in energy metabolism are necessary.

The increased detection of ANKRD5 in the midpiece region of the sperm axoneme does not necessarily indicate its localization in mitochondria. Immunofluorescence signals of multiple axonemal Nexin-Dynein Regulatory Complex (N-DRC) components (e.g., TCTE1, DRC1, CCDC65, DRC3, GAS8, and DRC7) are also non-uniformly distributed along the entire flagellum[1]. Similar localization patterns are observed in other structural components, such as radial spoke protein NME5[2] and outer dynein arm protein DNAH5[3]. Furthermore, mitochondria are membrane-bound organelles, and ANKRD5 predominantly resides in the SDS-soluble fraction under varying lysis conditions, confirming its association with the axoneme rather than mitochondria. Thus, the spatial distribution of ANKRD5 does not support a functional role in mitochondria. Importantly, we validated intact mitochondrial function through measurements of reactive oxygen species (ROS) levels (Figure S5C, D), ATP content (Figure 6E), and mitochondrial membrane potential (Figure S5A, B).

Proteomic analysis of Ankrd5-deficient sperm revealed disruptions in the glycolysis pathway. While these changes do not appear to affect ATP production, the mechanism by which these disruptions impact sperm motility remains unclear. Further investigation into how glycolysis pathway alterations contribute to impaired motility is warranted.

We appreciate the reviewer's careful consideration of our proteomic data. However, our Gene Set Enrichment Analysis (GSEA) of glycolysis/gluconeogenesis pathways showed no significant enrichment (p-value=0.089, NES=0.708; Fig.6D), which does not meet the statistical thresholds for biological significance (|NES|>1, pvalue<0.05). This observation is further corroborated by our direct ATP measurements showing no difference between genotypes (Fig.6E). We agree that further studies on metabolic regulation could be valuable, but current evidence does not support glycolysis disruption as a primary mechanism for the motility defects observed in Ankrd5-null sperm. This misinterpretation likely arose from the reviewer's overinterpretation of non-significant proteomic trends. We request that this specific claim be excluded from the assessment to avoid misleading readers.

Weaknesses:

Cryo-ET Data Limitations: The structural analysis of the DMT lacks clarity on how ANKRD5 influences NDRC or RS3. The low quality of RS3 data hinders the interpretation of ANKRD5's impact on axoneme structure.

We tried to further calculate the DMT at 96nm period using the present data to analyze the effect of ANKRD5 deletion on RS and N-DRC, however, due to the heterogeneity of the data, we were only able to obtain DMT at 8nm period (we have added a figure in the supplementary material for presentation). And in the process of obtaining DMT with a period of 8nm, we throw away about 90% of the particles (some are misaligned particles, some are deformed particles). Although we were not able to obtain the structure of 96nm repeats DMT, we noticed the enhanced heterogeneity of DMT caused by ANKRD5 knockout, as shown by the 3D classification and other results of the new supplementary images (Fig. S9), and the graphic description was added in the original article.

We have changed the description in the text: "During particle picking of DMT fibers, we observed that transverse sections of axonemal DMT particles from ANKRD5-KO sperm differ markedly from those in WT sperm. Although both A- and B-tubes were visible in both samples, the DMTs in ANKRD5-KO sperm showed a more irregular profile. In WT sperm, DMTs typically appeared circular, whereas ANKRD5-KO DMTs seemed to be extruded as polygonal. (Fig. S9B,D). Notably, ANKRD5-KO DMTs seemed partially open at the junction between the A- and B-tubes (Fig. S9B,D).

During the STA process, many particles were misaligned or deformed in the classification results, revealing various degrees of deformation—particularly affecting the B-tube (Fig. S9E). We could retain only ~10% of the DMT particles to obtain the final density map for ANKRD5-KO sperm (Fig. S9E), whereas ~70% were usable in WT dataset as reported previously [59]. The mutant DMT density map also displayed roughness at its periphery, indicating substantial structural heterogeneity (Fig. S9E). Even after discarding a large fraction of deformed particles, the final density map still showed evident artifacts, implying that although the mutant DMT preserves the fundamental features of both tubes, its shape is highly heterogeneous (Fig. S9E). Furthermore, attempts to classify the 96-nm repeats did not yield a clear density for radial spokes (RSs) (Fig. S9F), indicating that ANKRD5 deficiency may affect the stability of other accessory structures, such as RSs [23，24，25]. In the raw tomograms, RSs in ANKRD5-KO sperm appeared less regularly arranged than those in WT (Fig. S9A and C).

Most recently, following the submission of this work, ANKRD5 was reported to localize at the head of the N-DRC, simultaneously binding DRC11, DRC7, DRC4, and DRC5 [46]. This structural insight agrees with our in vitro findings that ANKRD5 interacts with DRC4 and DRC5 (Fig. 8C-F). However, that study used isolated and purified DMT samples, leaving the precise positioning of ANKRD5 between adjacent axonemal DMTs unconfirmed. We therefore fitted the published structure (PDB entry: 9FQR) into the in situ DMT structure of mouse sperm 96-nm repeats (EMD-27444), revealing that ANKRD5 lies a mere ~3 nm from the adjacent DMT (Fig. 8G). Notably, the N-DRC is often likened to a "car bumper", buffering two neighboring DMTs during vigorous axonemal motion. Given the extensive DMT deformation observed in our cryo-ET data (Fig. S9E), we propose that ANKRD5 contributes to this buffering function at the N-DRC. The loss of ANKRD5 may weaken the "bumper" effect and consequently increase structural damage to adjacent DMTs under intense conditions, while also compromising the stability of associated DMT accessory structures [19，46，60]."

Figure S9. The states of DMT particles in sperm of Ankrd5-KO mouse. (A) and (C) Tomogram slices of WT and Ankrd5-KO in Dynamo (The data for WT mouse sperm was EMPIARC-200007). DMT and RS are marked with white dashed lines and white arrows, respectively. (B) and (D) Comparison of DMT particle states between WT and Ankrd5-KO in Dynamo. The visual angles of the DMT particles shown in (B) and (D) show that the DMT fibers within the white box in (A) and (B) are divided equally into 10 slices along the direction of the white arrow, respectively. The DMT particle shapes of WT and Ankrd5-KO are marked by white dashed lines on the right of (B) and (D). The white arrow in (D) identifies the junction of A-tube and B-tube that is suspected to be disconnected. (E) Deformed particles discarded in 3D classification and final aligned DMT artifacts. (F) 3D classification of attempted RS locations.

Although the loss of ANKRD5 did not affect the density of DMT itself in A Tube and B Tube, we found that DMT particles were not smooth in the original tomogram. We speculate that the loss of ANKRD5, a component of the N-DRC that is close to the neighboring DMT, may reduce the interaction between N-DRC and the neighboring DMT, resulting in uneven force on the axoneme during sperm swimming, which may limit our ability to obtain the average structure of the more dynamic components (RS, N-DRC, ODA, IDA). Therefore, when trying to classify 96nm repeat DMT, we could only see the density of suspected RS3 and RS2, but it was difficult to obtain the complete 96nm repeat DMT density, so that we could not further analyze the effect of ANKRD5 deletion on RS and N-DRC. To test this conjecture, we further classified the density of suspected RS3, and the results obtained exhibited a variety of mixed states (which have been added to the supplementary material). To avoid confusion, we have already removed the discussion of RS3 and the related images from the original text.

The cryo-ET data on the internal structure of the DMT seems to have limited relevance to the N-DRC complex. Additionally, the quality of the RS3 data appears suboptimal, making it difficult to understand how the absence of ANKRD5 influences RS3. Further refinement of the data or alternative approaches may be needed to address this question.

Thank you very much for your suggestions. For the 96 nm periodic DMT, we have conducted multiple rounds of classification, including applying different masks at the positions of ODA, RS, and DMT. We have also tried classifying with both a single reference and multiple references. However, we were unable to obtain a suitable 96 nm periodic DMT. Regarding the heterogeneity of the particles, we have added a discussion in the manuscript. Following your advice, we have reanalyzed the data, but unfortunately, we still could not further optimize the experimental results.

In the process of obtaining the 8 nm periodic DMT, we discarded approximately 90 percent of the particles through multiple rounds of classification and alignment, in order to obtain high-quality 8 nm periodic DMT. We classified the remaining particles and found that the densities of RS3 and RS2 were not in their normal states. RS3 might be a mixture of different states of RS3, which makes it difficult for us to further discuss the effects of ANKRD5 on RS3.

To avoid confusion, we have already removed the discussion of RS3 and the related images from the original text.

Regarding the effects of ANKRD5 deficiency, we speculate that as the head of the N-DRC, its absence might affect the interaction between the N-DRC and the adjacent DMT, thereby influencing the forces experienced by the DMT during sperm movement. The uneven and irregular forces on the nine pairs of DMTs do not affect the structure of the A and B tubes of the DMT itself, but result in some heterogeneity in the peripheral microtubule parts of the DMT particles. We have added a discussion on these hypotheses in the manuscript. In addition, our 3D classification results demonstrate the structural heterogeneity of DMT caused by ANKRD5 knockdown. We have changed the description in the text:"During particle picking of DMT fibers, we observed that transverse sections of axonemal DMT particles from ANKRD5-KO sperm differ markedly from those in WT sperm. Although both A- and B-tubes were visible in both samples, the DMTs in ANKRD5-KO sperm showed a more irregular profile. In WT sperm, DMTs typically appeared circular, whereas ANKRD5-KO DMTs seemed to be extruded as polygonal. (Fig. S9B,D). Notably, ANKRD5-KO DMTs seemed partially open at the junction between the A- and B-tubes (Fig. S9B,D).

During the STA process, many particles were misaligned or deformed in the classification results, revealing various degrees of deformation—particularly affecting the B-tube (Figure 9, Fig. S9E). We could retain only ~10% of the DMT particles to obtain the final density map for ANKRD5-KO sperm (Fig. S9E), whereas ~70% were usable in WT dataset as reported previously [59]. The mutant DMT density map also displayed roughness at its periphery, indicating substantial structural heterogeneity (Fig. S9E). Even after discarding a large fraction of deformed particles, the final density map still showed evident artifacts, implying that although the mutant DMT preserves the fundamental features of both tubes, its shape is highly heterogeneous (Fig. S9E). Furthermore, attempts to classify the 96-nm repeats did not yield a clear density for radial spokes (RSs) (Fig. S9F), indicating that ANKRD5 deficiency may affect the stability of other accessory structures, such as RSs [23，24，25]. In the raw tomograms, RSs in ANKRD5-KO sperm appeared less regularly arranged than those in WT (Fig. S9A and C).

Most recently, following the submission of this work, ANKRD5 was reported to localize at the head of the N-DRC, simultaneously binding DRC11, DRC7, DRC4, and DRC5 [46]. This structural insight agrees with our in vitro findings that ANKRD5 interacts with DRC4 and DRC5 (Fig. 8C-F). However, that study used isolated and purified DMT samples, leaving the precise positioning of ANKRD5 between adjacent axonemal DMTs unconfirmed. We therefore fitted the published structure (PDB entry: 9FQR) into the in situ DMT structure of mouse sperm 96-nm repeats (EMD-27444), revealing that ANKRD5 lies a mere ~3 nm from the adjacent DMT (Fig. 8G). Notably, the N-DRC is often likened to a "car bumper", buffering two neighboring DMTs during vigorous axonemal motion. Given the extensive DMT deformation observed in our cryo-ET data (Fig. S9E), we propose that ANKRD5 contributes to this buffering function at the N-DRC. The loss of ANKRD5 may weaken the "bumper" effect and consequently increase structural damage to adjacent DMTs under intense conditions, while also compromising the stability of associated DMT accessory structures [19，46，60]."

To further enhance the readability of our manuscript, we created a Graphic Abstract to visually illustrate the biological functions of ANKRD5. The figure is placed immediately after the Abstract section and has been designated as Figure 9.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Sperm motility is crucial for male reproduction and relies on the structural integrity of the sperm axoneme, which has a “9+2” microtubule configuration. This structure includes nine outer microtubule doublets that house various macromolecular complexes. The nexin-dynein regulatory complex (N-DRC) forms a crossbridge between the outer microtubule doublets, stabilizing them and facilitates sperm tail bending. Our investigation of ANKRD5, which is highly expressed in the sperm axoneme, reveals its interaction with TCTE1 and DRC4/GAS8, both key components of the N-DRC. The components of the N-DRC are often vital for sperm motility. ANKRD5-/- mice exhibited reduced sperm motility and male infertility; however transmission electron microscopy and cryoelectron tomography showed no significant alterations in microtubule doublets. Moreover, ANKRD5 deficiency did not affect ATP levels, and its interactions with TCTE1 and DRC4/GAS8 were found to be independent of calcium regulation. These findings establish that ANKRD5 is critical for maintaining axoneme stability, which is important for sperm motility.

Revised version: Sperm motility is crucial for male reproduction and relies on the structural integrity of the sperm axoneme, which has a “9+2” microtubule configuration. This structure includes nine outer doublet microtubules (DMTs) that house various macromolecular complexes. The nexin-dynein regulatory complex (N-DRC) forms a crossbridge between adjacent DMTs, stabilizing them and facilitates sperm tail bending. Our study of ANKRD5, which is highly expressed in the sperm axoneme, reveals its interaction with DRC5/TCTE1 and DRC4/GAS8, both critical components of the N-DRC, and these interactions were found to be independent of calcium regulation. ANKRD5-/- mice exhibited reduced sperm motility and male infertility. Cryo-electron tomography analysis reveals typical “9+2” axoneme and intact DMT structures in Ankrd5-/- mouse sperm, but the DMTs displayed significant morphological variations and greater structural heterogeneity. Furthermore, ANKRD5 deficiency did not affect ATP level, ROS levels, or mitochondrial membrane potential. These findings suggest that ANKRD5 may weaken the N-DRC’s “car bumper” role, reducing the buffering effect between adjacent DMTs and thereby destabilizing axoneme structures during intense axoneme motility.

Introduction

Earlier version: The interaction between sperm and egg, leading to embryo formation, is fundamental for life reproduction and species continuity (1). Male infertility impacts 8%-12% of the global male population, with sperm motility defects contributing to 40%-50% of these cases (2, 3) Fertilization relies on successful spermatogenesis and normal sperm motility (4), which occurs in the testes. Mammalian sperm gain motility and fertilization capabilities when they pass through the epididymis (5). The maturation of sperm in the epididymis is necessary for the production of fertile sperm. Asthenospermia caused by poor sperm motility is the main symptom of clinical infertility, yet its precise pathogenesis remains largely unclear (6). Individuals with poorly motile or immobile sperm are typically infertile unless advanced assisted reproductive techniques, such as gamete intrafallopian transfer (GIFT), in vitro fertilization (IVF), or intracytoplasmic sperm injection (ICSI), are employed (7). However, these common assisted reproductive technologies may pass genetic defects to offspring. A more comprehensive understanding of the molecular mechanisms underlying sperm motility could lead to more effective clinical interventions for hypofertility caused by asthenospermia. Rather than bypassing the issue with ICSI, infertility from poor sperm motility could potentially be treated or even cured through stimulation of specific signaling pathways or gene therapy.

Revised version: The interaction between sperm and egg, leading to embryo formation, is fundamental for life reproduction and species continuity[1]. Male infertility impacts 8%-12% of the global male population, with sperm motility defects contributing to 40%-50% of these cases[2, 3] Fertilization relies on successful spermatogenesis and normal sperm motility[4]. Mammalian sperm gain motility and fertilization capabilities when they pass through the epididymis[5]. The maturation of sperm in the epididymis is necessary for the production of fertile sperm. Asthenospermia caused by poor sperm motility is the main symptom of clinical infertility, yet its precise pathogenesis remains largely unclear[6]. Individuals with poorly motile or immobile sperm are typically infertile unless advanced assisted reproductive techniques, such as gamete intrafallopian transfer (GIFT), in vitro fertilization (IVF), or intracytoplasmic sperm injection (ICSI), are employed[7]. However, these common assisted reproductive technologies may pass genetic defects to offspring. A more comprehensive understanding of the molecular mechanisms underlying sperm motility could lead to more effective clinical interventions for hypofertility caused by asthenospermia. Rather than bypassing the issue with ICSI, infertility from poor sperm motility could potentially be treated or even cured through stimulation of specific signaling pathways or gene therapy.

Earlier version: Sperm motility is driven by flagellar beating. The flagellum consists of three main parts: the midpiece, principal piece, and endpiece (8). These three parts share a common central axoneme composed of approximately 250 proteins that are major components of the flagellum (9). The axoneme exhibits a 9+2 microtubule doublet structure, with the N-DRC forming cross-bridges between the outer microtubule doublets (10). Both the structure and molecular composition of N-DRC are evolutionarily conserved and play a central role in regulating sperm motility (11–13). The N-DRC is a large, complex structure with a molecular weight of approximately 1.5 MDa and can be divided into two regions: the linker and the base plate (12–14). The N-DRC is a large, complex structure with a molecular weight of approximately 1.5 MDa and can be divided into two regions: the linker and the base plate (14–16). Additionally, the N-DRC interacts with the outer dynein arms (ODA) via outer-inner dynein (OID) linkers, functioning as a regulatory component for both ODA and inner dynein arms (IDA) (17).

Revised version: Sperm motility is driven by flagellar beating. The flagellum consists of three main parts: the midpiece, principal piece, and endpiece[8]. These three parts share a common central axoneme composed of approximately 250 proteins that are major components of the flagellum[9]. The axoneme features a 9+2 architecture, comprising nine doublet microtubules (DMTs) encircling a central pair of singlet microtubules, with the N-DRC forming connections between adjacent DMTs[10]. Both the structure and molecular composition of N-DRC are evolutionarily conserved and play a central role in regulating sperm motility[11, 12, 13]. The N-DRC is a large, complex structure with a molecular weight of approximately 1.5 MDa and can be divided into two regions: the linker and the base plate[12, 13, 14]. The N-DRC is a large, complex structure with a molecular weight of approximately 1.5 MDa and can be divided into two regions: the linker and the base plate[14, 15, 16]. Additionally, the N-DRC interacts with the outer dynein arms (ODA) via outer-inner dynein (OID) linkers, functioning as a regulatory component for both ODA and inner dynein arms (IDA)[17].

Earlier version: It was previously believed that N-DRC comprised 11 protein components (13, 18). However, integrated modeling combined with in situ cross-linking mass spectrometry revealed the existence of CCDC153 (DRC12), which is cross-linked with DRC1 (19). In situ cryoelectron tomography (cryo-ET) has further advanced understanding of the N-DRC architecture in Chlamydomonas, demonstrating that DRC1, DRC2/CCDC65, and DRC4/GAS8 constitute its core framework (16). Additionally, proteins DRC3/5/6/7/8/11 associate with this framework and engage with other axonemal complexes (20). Biochemical experiments corroborate these findings and support this model. The integrity of the N-DRC structure is crucial for flagellar movements, and mutations of N-DRC subunits may lead to defects in flagellar movement (15, 21–23). The N-DRC functions between the doublet microtubules (DMTs) to convert sliding into axonemal bending motion by restricting the relative sliding of outer microtubule doublets (24–26). Additionally, the formation of the N-DRC necessitates the coordination among various subunits. Mutations in three specific subunits—DRC1, DRC2/CCDC65, and DRC4/GAS8—have been previously associated with human ciliary motility disorders leading to primary ciliary dyskinesia (PCD) (12, 27). Knockout or mutation of Lrrc48 (Drc3) (28), Tcte1 (Drc5) (29), Drc7 (30), and Iqcg (Drc9) (31) in mouse has been shown to cause male infertility. Our findings indicate that ANKRD5 (Ankyrin repeat domain 5, ANKRD5, ANK5, ANKEF1) could interact with N-DRC structure, serving as an auxiliary element to facilitate collaboration among DRC members. The absence of ANKRD5 results in diminished sperm motility and subsequent male infertility.

Revised version: It was previously believed that N-DRC comprised 11 protein components[13, 18]. However, a new component CCDC153 (DRC12) was found to interact with DRC1[19]. In situ cryoelectron tomography (cryo-ET) has significantly advanced understanding of the N-DRC architecture in Chlamydomonas, demonstrating that DRC1, DRC2/CCDC65, and DRC4/GAS8 constitute its core framework[16], while proteins DRC3/5/6/7/8/11 associate with this framework and engage with other axonemal complexes[20]. Biochemical experiments corroborate these findings and validate this structural model[12, 21, 22]. The N-DRC functions between the DMTs to convert sliding into axonemal bending motion by restricting the relative sliding of outer microtubule doublets[23, 24, 25]. Mutations of N-DRC subunits demonstrate that the structural integrity of the N-DRC is crucial for flagellar movements. Mutations in DRC1, DRC2/CCDC65, and DRC4/GAS8 are linked to ciliary motility disorders, causing primary ciliary dyskinesia (PCD)[12, 26]. Biallelic truncating mutations in DRC1 induce multiple morphological abnormalities of sperm flagella (MMAF), including outer DMT disassembly, mitochondrial sheath disorganization, and incomplete axonemal structures in human sperm[22, 27, 28]. Similarly, CCDC65 loss disrupts N-DRC stability, leading to disorganized axonemes, global microtubule dissociation, and complete asthenozoospermia[12, 29]. Homozygous frameshift mutations in DRC3 impair N-DRC assembly and intraflagellar transport (IFT), resulting in severe motility defects despite normal sperm morphology[30, 31]. TCTE1 knockout mice maintain normal sperm axoneme structure but show impaired glycolysis, leading to reduced ATP levels, lower sperm motility, and male infertility[32]. Both Drc7 and Iqcg (Drc9) knockout mice exhibit disrupted ‘9+2’ axonemal architecture, sperm immotility, and male infertility[21, 33]. Drc7 knockout sperm also display head deformities and shortened tails[21]. While N-DRC is critical for sperm motility, but the existence of additional regulators that coordinate its function remains unclear. Our findings indicate that ANKRD5 (Ankyrin repeat domain 5; also known as ANK5 or ANKEF1) interacts with N-DRC structure, serving as an auxiliary element to facilitate collaboration among DRC members. The absence of ANKRD5 results in diminished sperm motility and consequent male infertility.

Results

Ankrd5 is critical for male fertility

Earlier version: The auxiliary factors required for the N-DRC structure to function effectively have not been fully elucidated. As an important structural component of the axoneme, the functionality of the N-DRC relies on protein interactions. The ANK domain is one of the crucial domains for protein interactions. Based on an analysis of the NCBI database and relevant single-cell sequencing data, Ankrd5 expression is predominantly localized to the male reproductive system, with a relatively high expression specificity (32). However, the impact of ANKRD5 on the functionality of the N-DRC remains unclear. In mice, ANKRD5 is a protein comprising 775 amino acids with a molecular weight of 86.9 kDa. Comparative analysis of the ANKRD5 protein sequence across various species revealed that it is relatively conserved (Fig. S1A). Sequence alignment using Clustal Omega demonstrated an 86% similarity between ANKRD5 sequences in mice and humans (Fig. S1B), suggesting that studying ANKRD5 in mice has significant translational relevance for human applications. We further confirmed this tissue-specific expression in mice through qPCR, demonstrating that Ankrd5 was highly expressed in the testes, with no detectable levels in the brain, liver, spleen, kidney, ovary, intestine, or stomach (Fig. 1A). Further investigation into the temporal expression of Ankrd5 in mouse testes indicated that its expression begins on postnatal day 21 and reaches peak levels by postnatal day 35 (Fig. 1B), which coincides with the emergence of mature sperm.

Revised version: The auxiliary factors required for the N-DRC structure to function effectively have not been fully elucidated. As an important structural component of the axoneme, the functionality of the N-DRC relies on protein interactions. The ANK domain is one of the crucial domains for protein interactions. Based on an analysis of the NCBI database and relevant single-cell sequencing data, Ankrd5 expression is predominantly localized to the male reproductive system, with a relatively high expression specificity[34]. However, the impact of ANKRD5 on the functionality of the N-DRC remains unclear. In mice, ANKRD5 is a protein comprising 775 amino acids with a molecular weight of 86.9 kDa. Comparative analysis of the ANKRD5 protein sequence across various species revealed that it is relatively conserved (Fig. S1A). Sequence alignment using Clustal Omega demonstrated an 86% similarity between ANKRD5 sequences in mice and humans (Fig. S1B), suggesting that studying ANKRD5 in mice has significant translational relevance for human applications. We further confirmed this tissue-specific expression in mice through qPCR, demonstrating that Ankrd5 was highly expressed in the testes, with no detectable levels in the brain, liver, spleen, kidney, ovary, intestine, or stomach (Fig. 1A). Further investigation into the temporal expression of Ankrd5 in mouse testes indicated that its expression begins on postnatal day 21 and reaches peak levels by postnatal day 35 (Fig. 1B), which coincides with the emergence of mature sperm.

Added in the revised version: Ankrd5 is critical for male reproductive function.

Earlier version: Ankrd5 is critical for male reproductive function. (A and B) Relative expression of Ankrd5 mRNA in different tissues of adult mouse and testes at various postnatal days. The Ankrd5 mRNA expression levels were normalized to the expression of Gapdh mRNA (n=3). (C) CRISPR/Cas9 targeting scheme of mouse Ankrd5 and genotyping of Ankrd5 KO mouse. Ankrd5-WT-F + Ankrd5-screen-R ( for WT) and Ankrd5-screen-F + Ankrd5-screen-R (for KO). nc, negative control (ddH2O). (D and E) Sperm count and percentage of normal sperm of cauda epididymal from control and Ankrd5 KO mouse (n=5). (F) Testis to body weight ratio of adult control and Ankrd5 KO mouse (n=?). (G) Hematoxylin and eosin ( H&E) staining of mouse testis and epididymis. Coomassie Brilliant Blue R-250 staining of spermatozoa from control and Ankrd5 KO male mouse. No significant abnormality was found in Ankrd5 KO male mouse. No overt abnormalities were found in Ankrd5 KO mouse. P, pachytene; ES, elongated sperm; RS, round sperm; SG, spermatogonia; ST, Sertoli cell. All values in this figure are shown as the mean± SE.

Revised version: (A and B) Relative expression of Ankrd5 mRNA in different tissues of adult mouse and testes at various postnatal days. The Ankrd5 mRNA expression levels were normalized to the expression of Gapdh mRNA (n=3). (C) CRISPR/Cas9 targeting scheme of mouse Ankrd5 and genotyping of Ankrd5 KO mouse. Ankrd5-WT-F + Ankrd5-screen-R (for WT) and Ankrd5-screen-F + Ankrd5-screen-R (for KO). nc, negative control (ddH2O). (D and E) Sperm count and percentage of normal sperm of cauda epididymal from control and Ankrd5 KO mouse (n=5). (F) Testis to body weight ratio of adult control and Ankrd5 KO mouse (n=7). (G) Hematoxylin and eosin (H&E) staining of mouse testis and epididymis. Coomassie Brilliant Blue R-250 staining of spermatozoa from control and Ankrd5 KO male mouse. No significant abnormality was found in Ankrd5 KO male mouse. No overt abnormalities were found in Ankrd5 KO mouse. P, pachytene; ES, elongated sperm; RS, round sperm; SG, spermatogonia; ST, Sertoli cell. All values in this figure are shown as the mean ± SE.

Earlier version: To explore the biological function of Ankrd5 in vivo, we generated an Ankrd5 knockout mouse line (Ankrd5−/−) on a C57BL/6J background by targeting exons 4 and 7 using CRISPR/Cas9. The knockout was confirmed at the gene level (Fig. 1C). Fertility tests showed that Ankrd5−/− males were able to mate with wild-type females but did not produce any offspring (Table 1). These results indicated that Ankrd5 is essential for male fertility. Since there were no significant differences in litter size (Table 1) and spermatogenesis between WT and het male mice (Fig. 1 D and E), we used het as the experimental control. No significant differences were observed in spermatogenesis between Ankrd5−/− and control mice (Fig. 1 DE and G). The testis-to-body weight ratio and the morphology of the reproductive system in Ankrd5−/− mice were also normal (Fig. 1F and S1C-D). Histological analysis of hematoxylin-eosin-stained paraffin sections further confirmed that the structure and cellular composition of the seminiferous tubules were comparable between Ankrd5−/− and control mice (Fig. 1G). These results indicated that the infertility phenotype in Ankrd5−/− mice was not a result of defective spermatogenesis.

Revised version: To explore the biological function of Ankrd5 in vivo, we generated an Ankrd5 knockout mouse line (Ankrd5−/–) on a C57BL/6J background by targeting exons 4 and 7 using CRISPR/Cas9. The knockout was confirmed at the gene level (Fig. 1C). Fertility tests showed that Ankrd5−/– males were able to mate with wild-type females but did not produce any offspring (Table 1). These results indicated that Ankrd5 is essential for male fertility. Since there were no significant differences in litter size (Table 1) and spermatogenesis between WT and Ankrd5+/– male mice (Fig. 1 D and E), we used Ankrd5+/– as the experimental control. No significant differences were observed in spermatogenesis between Ankrd5−/– and control mice (Fig. 1 DE and G). The testis-to-body weight ratio and the morphology of the reproductive system in Ankrd5−/– mice were also normal (Fig. 1F and S1C-D). Histological analysis of hematoxylin-eosin-stained paraffin sections further confirmed that the structure and cellular composition of the seminiferous tubules were comparable between Ankrd5−/– and control mice (Fig. 1G). These results indicated that the infertility phenotype in Ankrd5−/– mice was not a result of defective spermatogenesis.

Due to decreased sperm motility, Ankrd5 null sperm were unable to penetrate the zona pellucida

Earlier version: To investigate the cause of infertility in Ankrd5 knockout mice, we performed in vitro fertilization (IVF) experiments. control sperm successfully fertilized both cumulus-intact eggs and cumulus-free eggs (Fig. 2A and B). In contrast, Ankrd5 null sperm failed to fertilize the cumulus-intact eggs despite normal binding to the zona pellucida (ZP) (Fig. 2 A and D). However, Ankrd5 null sperm were able to fertilize zona pellucida-free eggs and develop to blastocyst (Fig. 2C). These results indicate that the infertility observed in Ankrd5 null sperm is primarily due to the sperm’s inability to penetrate the zona pellucida. Sperm penetration through the cumulus-oocyte complex (COC) is influenced by two key factors: the acrosome reaction and sperm motility (33). Notably, Ankrd5 null sperm exhibit a normal acrosome reaction when stimulated with calcium ionophore A23187 (Fig. 2 E and F), suggesting that the reproductive defects are not attributable to an acrosome reaction deficiency.

Revised version: To investigate the cause of infertility in Ankrd5 knockout mice, we performed in vitro fertilization (IVF) experiments. Control sperm successfully fertilized both cumulus-intact eggs and cumulus-free eggs (Fig. 2A and B). In contrast, Ankrd5 null sperm failed to fertilize the cumulus-intact eggs despite normal binding to the zona pellucida (ZP) (Fig. 2 A and D). However, Ankrd5 null sperm were able to fertilize zona pellucida-free eggs and develop to blastocyst (Fig. 2C). These results indicate that the infertility observed in Ankrd5 null sperm is primarily due to the sperm’s inability to penetrate the zona pellucida. Sperm penetration through the cumulus-oocyte complex (COC) is influenced by two key factors: the acrosome reaction and sperm motility[35]. Notably, Ankrd5 null sperm exhibit a normal acrosome reaction when stimulated with calcium ionophore A23187 (Fig. 2 E and F), suggesting that the reproductive defects are not attributable to an acrosome reaction deficiency.

Added in the revised version: Evaluation of in vitro fertilization capacity of Ankrd5 KO sperm.

Earlier version: Evaluation of in vitro fertilization capacity of Ankrd5 KO sperm. (A-C) Fertilization rate of IVF using control and Ankrd5 KO spermatozoa. Three types of oocytes (cumulus-intact, cumulus-free, and zona pellucida-free) were used for IVF. (D) Egg observation after IVF. After 4 hours of incubation, both of control and Ankrd5 KO sperm could penetrate cumulus oophorus as indicated by the red arrow and have the ability to bind to the zona pellucida. (E and F) Sperm were incubated in capacitation medium treated with A23187 (dissolved in DMSO) and DMSO (dissolvent control group) and stained with Coomassie Brilliant Blue R-250. Black arrow indicates the intact or disappeared acrosome; Values represent mean ± SE (n=3).

Revised version: (A-C) Fertilization rate of IVF using control and Ankrd5 KO spermatozoa. Three types of oocytes (cumulus-intact, cumulus-free, and zona pellucida-free) were used for IVF. (D) Egg observation after IVF. After 4 hours of incubation, both of control and Ankrd5 KO sperm could penetrate cumulus oophorus as indicated by the red arrow and have the ability to bind to the zona pellucida. (E and F) Sperm were incubated in capacitation medium treated with A23187 (dissolved in DMSO) and DMSO (dissolvent control group) and stained with Coomassie Brilliant Blue R-250. Black arrow indicates the intact or disappeared acrosome; Values represent mean ± SE (n=3).

Sperm motility is critical for penetrating the cumulus oophorus and zona pellucida, which are softened by the acrosome reaction. Computer-assisted sperm analysis (CASA) revealed significant reductions in curvilinear velocity (VCL), average path velocity (VAP), and straight line velocity (VSL) in Ankrd5 null sperm compared to controls (Fig. 3A). Additionally, forward progression parameters such as straightness (STR) and linearity (LIN) were also diminished (Fig. 3A). Sperm were categorized into four motility grades: rapid, medium, slow, and static. Significant differences were observed in the distribution of rapid, slow, and static grades, while no significant difference was noted in the medium grade (Fig. 3B). The proportions of total and progressive motility decreased with the loss of Ankrd5 function (Fig. 3C). Furthermore, the tracking of sperm trajectories over time revealed that Ankr5- null sperm exhibited lower motility compared to controls (Fig. 3D). Moreover, sperm migration experiments revealed a significantly lower presence of Ankrd5 null sperm in the female uterus and oviducts six hours after mating (Fig. 3E-G), further supporting the observed motility defects. These findings suggest that the reduced motility of Ankrd5 null sperm compromised their ability to penetrate the outer layers of the egg.

Added in the revised version: Sperm motility of Ankrd5 KO male mouse.

Earlier version: Sperm motility of Ankrd5 KO male mouse. (A) Average path velocity (VAP), straight line velocity (VSL), curvilinear velocity (VCL), amplitude of lateral head displacement (ALH), beat cross frequency (BCF), straightness (STR), and linearity (LIN) of sperm from control and Ankrd5 KO mouse. **P < 0.01, ***P < 0.001, Student’s t test; error bars represent SE (n = 3). (B) Proportions of sperm at different velocity levels in control and Ankrd5 KO mouse. **P < 0.01, ***P < 0.001, Student’s t test; error bars represent SE (n = 3). (C) Knockout mouse had lower motile sperm (total motor capacity) and progressive motile sperm (progressive motor capacity) than control. **P < 0.01, ***P < 0.001, Student’s t test; error bars represent SE (n = 3). (D) Trajectories of sperm per second. The meanings of different colors are shown in the graph. (E) Impaired migration of Ankrd5 KO sperm from uterus into oviducts. The black arrow indicates sperm. (F and G) Numbers of sperm from control and Ankrd5 KO mouse in female uterus and oviducts after mating. **P < 0.01, ***P <0.001, Student’s t test; error bars represent SE (n = 3).

Revised version: (A) Average path velocity (VAP), straight line velocity (VSL), curvilinear velocity (VCL), amplitude of lateral head displacement (ALH), beat cross frequency (BCF), straightness (STR), and linearity (LIN) of sperm from control and Ankrd5 KO mouse. **P < 0.01, ***P < 0.001, Student’s t test; error bars represent SE (n = 3). (B) Proportions of sperm at different velocity levels in control and Ankrd5 KO mouse. **P < 0.01, ***P < 0.001, Student’s t test; error bars represent SE (n = 3). (C) Knockout mouse had lower motile sperm (total motor capacity) and progressive motile sperm (progressive motor capacity) than control. **P < 0.01, ***P < 0.001, Student’s t test; error bars represent SE (n = 3). (D) Trajectories of sperm per second. The meanings of different colors are shown in the graph. (E) Impaired migration of Ankrd5 KO sperm from uterus into oviducts. The black arrow indicates sperm. (F and G) Numbers of sperm from control and Ankrd5 KO mouse in female uterus and oviducts after mating. **P < 0.01, ***P <0.001, Student’s t test; error bars represent SE (n = 3).

Earlier version: ANKRD5 is primarily localized in the sperm axoneme

Revised version: ANKRD5 localized in the sperm axoneme and tracheal motile cilia

To explore the biological function of ANKRD5, we generated Ankrd5-FLAG mice by inserting a FLAG tag at the C-terminus using homologous recombination (Fig. S2A). Tissue sectioning and hematoxylin and eosin staining were performed to evaluate the epididymal structure and cellular composition in Ankrd5-FLAG male mice compared to controls (Fig. S2 B and C). The results revealed that the epididymis exhibited a normal structure, and the reproductive capabilities of Ankrd5-FLAG mice were intact, indicating that FLAG insertion did not impair the biological function of ANKRD5. Western blot analysis was employed to examine ANKRD5 protein expression across multiple organs and tissues, including liver, lung, kidney, heart, testis, spleen, brain, small intestine, skin, and sperm. The data revealed that ANKRD5 was predominantly expressed in mouse sperm (Fig. 4A), with lower expression levels in the testis. To further investigate ANKRD5 localization within sperm, we subjected sperm to repeated freeze-thaw cycles followed by density gradient centrifugation to separate the head and tail. Western blot analysis confirmed that ANKRD5 was primarily localized in the tail (Fig. 4B).

Added in the revised version: ANKRD5 is located in the midpiece of sperm axoneme.

Earlier version: ANKRD5 is located in the midpiece of sperm axoneme. (A) Immunoblotting analysis of various mouse tissues. GAPDH was used as the loading control. (B) Head and tail separation of mouse spermatozoa. ANKRD5-FLAG was detected in the tail fraction. PRM2 was used as a marker for sperm head. BASIGIN, AKAP4 and acetylated tubulin were detected as a marker for tails. (C) Immunofluorescence staining results of spermatozoa from wild-type and ANKRD5-FLAG mouse using anti-FLAG antibody (red: anti-FLAG signal; Hoechst: blue). (D) Fractionation of sperm proteins using different lysis buffers. ANKRD5-FLAG was found in the SDS-soluble fraction that contains axonemal proteins. BASIGIN, acetylated tubulin, and AKAP4 were detected as a marker for Triton-soluble, SDS-soluble, and SDS-resistant fractions, respectively. (E) Ultrastructure of sperm tails in control and Ankrd5 KO mouse. Cross-sections of midpiece, principal piece and end piece were observed using transmission electron microscopy.

Revised version: (A) Immunoblotting analysis of various mouse tissues. GAPDH was used as the loading control. (B) Head and tail separation of mouse spermatozoa. ANKRD5-FLAG was detected in the tail fraction. PRM2 was used as a marker for sperm head. BASIGIN, AKAP4 and acetylated tubulin were detected as a marker for tails. (C) Immunofluorescence staining results of spermatozoa from wild-type and ANKRD5-FLAG mouse using anti-FLAG antibody (red: anti-FLAG signal; Hoechst: blue). (D) Fractionation of sperm proteins using different lysis buffers. ANKRD5-FLAG was found in the SDS-soluble fraction that contains axonemal proteins. BASIGIN, acetylated tubulin, and AKAP4 were detected as a marker for Triton-soluble, SDS-soluble, and SDS-resistant fractions, respectively. (E) Ultrastructure of sperm tails in control and Ankrd5 KO mouse. Cross-sections of midpiece, principal piece and end piece were observed using transmission electron microscopy.

Earlier version: The sperm tail comprises three sections: midpiece, principal piece, and end piece. Immunofluorescence staining of adult mouse sperm revealed that ANKRD5 was primarily localized in the midpiece (Fig. 4C). As reported in previous studies, proteins from different regions of the sperm tail can be extracted with varying lysate intensities (29, 34, 35). To determine the precise localization of ANKRD5, we used specific markers, including BASIGIN, Acetylated Tubulin, and AKAP4, representing the Triton X-100 soluble, SDS soluble, and SDS resistant fractions, respectively. These three fractions contain, respectively, membrane-associated and cytosolic proteins, axonemal proteins, as well as fibrous sheath and outer dense fiber proteins (29, 34, 35). The results indicated that ANKRD5 was primarily located in the SDS soluble fraction, suggesting its association with the sperm axoneme (Fig. 4D). These findings suggest that the infertility associated with the reduced motility in Ankrd5 null sperm may be linked to impaired axoneme function in the sperm tail.

Revised version: The sperm tail comprises three sections: midpiece, principal piece, and end piece. Immunofluorescence staining of adult mouse sperm revealed that ANKRD5 was primarily localized in the midpiece (Fig. 4C). As reported in previous studies, proteins from different regions of the sperm tail can be extracted with varying lysate intensities[36, 37, 38]. To determine the precise localization of ANKRD5, we used specific markers, including BASIGIN, Acetylated Tubulin, and AKAP4, representing the Triton X-100 soluble, SDS soluble, and SDS resistant fractions, respectively. These three fractions contain, respectively, membrane-associated and cytosolic proteins, axonemal proteins, as well as fibrous sheath and outer dense fiber proteins[36, 37, 38]. The results indicated that ANKRD5 was primarily located in the SDS soluble fraction, suggesting its association with the sperm axoneme (Fig. 4D). These findings suggest that the infertility associated with the reduced motility in Ankrd5 null sperm may be linked to impaired axoneme function in the sperm tail. Given that mouse respiratory cilia also possess a canonical axonemal structure[39], our immunostaining analysis confirmed the expression of ANKRD5 in respiratory cilia (Fig. S3). However, Ankrd5-KO mice exhibited no significant survival defects or locomotor impairments beyond male infertility, suggesting a tissue-specific functional dependency of ANKRD5.

ANKRD5 interacts with axoneme N-DRC components

Earlier version: In order to explore the function of ANKRD5 in the axoneme, we explored the interactome of ANKRD5 using LC-MS (Fig. 5A). We identified the presence of several N-DRC components, including TCTE1, DRC3/LRRC48, DRC7/CCDC135, and DRC4/GAS8, with TCTE1 ranking highly in the mass spectrum (Fig. 5A). TEM analysis indicates that Tcte1 null sperm maintain a normal DMT structure yet exhibit reduced motility (29), a phenotype similarly observed in ANKRD5 knockout sperm. ANKRD5 was identified in the TCTE1 mass spectrum with a high ranking, but their interaction had not been experimentally validated before our study (29). We performed immunoprecipitation in 293T cells and validated protein interactions with 11 DRC components, except for DRC6, which failed to express (Fig. 5 B and C, Fig. S3 A and B). Our results confirmed that ANKRD5 interacts with TCTE1 and DRC4/GAS8, but no interaction with other DRC components (Fig. 5 B and C, Fig. S3 A and B). ANKRD5 consists of two ANK domains and one EF-hand domain. The ANK domain is classic mediator of protein interactions, and the EF-hand domain typically functions as a calcium sensor; however, approximately one-third of EF-hand domains lack calcium-binding capability (36). Through protein truncation experiments, we found that the ANK2 domain mediates the interaction between ANKRD5 and TCTE1 (Fig. 5E). But, the interaction between ANKRD5 and DRC4/GAS8 requires the presence of both the ANK1 and ANK2 domains (Fig. 5F). The EF-Hand domain is not necessary for the interaction between ANKRD5 and TCTE1 or DRC4/GAS8. Western blot results indicated that calcium ions do not affect the interaction between ANKRD5 and DRC4/GAS8 or TCTE1 (Fig. 5G).

Revised version: In order to explore the function of ANKRD5 in the axoneme, we explored the interactome of ANKRD5 using LC-MS (Fig. 5A). We identified the presence of several N-DRC components, including TCTE1, DRC3/LRRC48, DRC7/CCDC135, and DRC4/GAS8, with TCTE1 ranking highly in the mass spectrum (Fig. 5A). TEM analysis indicates that Tcte1 null sperm maintain a normal DMT structure yet exhibit reduced motility[36], a phenotype similarly observed in ANKRD5 knockout (ANKRD5-KO) sperm. ANKRD5 was identified in the TCTE1 mass spectrum with a high ranking, but their interaction had not been experimentally validated before our study[36]. We performed immunoprecipitation in 293T cells and validated protein interactions with 11 DRC components, except for DRC6, which failed to express (Fig. 5 B and C, Fig. S3 A and B). Our results confirmed that ANKRD5 interacts with TCTE1 and DRC4/GAS8, but no interaction with other DRC components (Fig. 5 B and C, Fig. S3 A and B). ANKRD5 consists of two ANK domains and one EF-hand domain. The ANK domain is classic mediator of protein interactions, and the EF-hand domain typically functions as a calcium sensor; however, approximately one-third of EF-hand domains lack calcium-binding capability[40]. Through protein truncation experiments, we found that the ANK2 domain mediates the interaction between ANKRD5 and TCTE1 (Fig. 5E). But, the interaction between ANKRD5 and DRC4/GAS8 requires the presence of both the ANK1 and ANK2 domains (Fig. 5F). The EF-Hand domain is not necessary for the interaction between ANKRD5 and TCTE1 or DRC4/GAS8. Western blot results indicated that calcium ions do not affect the interaction between ANKRD5 and DRC4/GAS8 or TCTE1 (Fig. 5G).

Added in the revised version: ANKRD5 is a component of N-DRC in sperm flagella.

Earlier version: ANKRD5 is a component of N-DRC in sperm flagella. (A) Identification of sperm proteins in LC-MS/MS analysis. Black star indicates N-DRC components. (B and C) Individual DRC components were coexpressed in HEK293T cells. Immunoprecipitation of ANKRD5-FLAG resulted in the co-precipitation of GAS8-MYC and TCTE1-MYC. Similarly, immunoprecipitation of GAS8-MYC and TCTE1-MYC also led to the co-precipitation of ANKRD5-FLAG. (D) Schematic of various truncated ANKRD5 vectors. FLAG-tag was linked posterior to the C-terminal of ANKRD5. Green and yellow boxes show the ANK domain and EF-Hand domain of ANKRD5, respectively. Light yellow boxes indicate FLAG tag. (E and F) The interaction between various truncated ANKRD5-FLAG and TCTE1-MYC or GAS8-MYC were confirmed by co-IP followed by WB analysis using anti-FLAG, and anti-MYC antibodies. (G) Effect of calcium ion and EDTA treatment on the interaction of ANKRD5 with GAS8 and TCTE1.

Revised version: (A) Identification of sperm proteins in LC-MS/MS analysis. Black star indicates N-DRC components. (B and C) Individual DRC components were coexpressed in HEK293T cells. Immunoprecipitation of ANKRD5-FLAG resulted in the co-precipitation of GAS8-MYC and TCTE1-MYC. Similarly, immunoprecipitation of GAS8-MYC and TCTE1-MYC also led to the co-precipitation of ANKRD5-FLAG. (D) Schematic of various truncated ANKRD5 vectors. FLAG-tag was linked posterior to the C-terminal of ANKRD5. Green and yellow boxes show the ANK domain and EF-Hand domain of ANKRD5, respectively. Light yellow boxes indicate FLAG tag. (E and F) The interaction between various truncated ANKRD5-FLAG and TCTE1-MYC or GAS8-MYC were confirmed by co-IP followed by WB analysis using anti-FLAG, and anti-MYC antibodies. (G) Effect of calcium ion and EDTA treatment on the interaction of ANKRD5 with GAS8 and TCTE1.

Added in the revised version: As for other proteins in the LC-MS results, KRT77 is a classic protein that maintains cytoskeletal stability. It may enhance the physical connection between the N-DRC and adjacent DMTs through interaction with ANKRD5. Recent studies indicate that ANKRD5, a newly identified component in the distal lobe of the N-DRC, has a positively charged surface, which may facilitate binding to glutamylated tubulin on adjacent DMTs[41]. Thus, KRT77 may also regulate its interaction with ANKRD5 via post-translational modifications (PTMs, e.g., phosphorylation), thereby strengthening sperm resistance to shear forces during flagellar movement. Rab family proteins participate in intraflagellar transport and membrane dynamics. RAB2A may promote targeted transport of ANKRD5 or other N-DRC components to axonemal assembly sites by recruiting vesicles, and its GTPase activity might link cellular signals to ANKRD5-mediated axoneme remodeling. However, the observed signals could be false positives due to nonspecific factors such as electrostatic adsorption, high-abundance protein interference, detergent-induced membrane disruption, or protein aggregation tendencies.

Earlier version: As is well known, energy metabolism is crucial for sperm motility. The TCTE1 null sperm showed a decrease in ATP levels (29). The loss of ANKRD5 does not change ATP levels in the sperm (Fig. 6E). This result indicates that ANKRD5 can interact with N-DRC components without impacting energy metabolism. The mitochondrial sheath is located in the midpiece of sperm, where mitochondria play a key role in energy production and serve as a significant source of reactive oxygen species (ROS) (37, 38). A balanced level of ROS is crucial for several essential sperm functions, including motility, capacitation, the acrosome reaction, fertilization, and hyperactivation (39–41). Mitochondrial membrane potential (MMP) is correlated with sperm motility (42). Decreased MMP (depolarization) indicates mitochondrial dysfunction, while increased MMP (hyperpolarization) leads to excess ROS. MMP in spermatozoa was assessed using TMRM staining and fluorescence imaging (Fig. S4 A and B). ROS levels were measured using DCFH-DA staining, followed by fluorescence imaging (Fig. S4 C and D). We found no significant differences between Ankrd5 null sperm and control. Thus, ANKRD5 deficiency did not significantly affect sperm ATP levels and mitochondrial function.

Revised version: As is well known, energy metabolism is crucial for sperm motility. The TCTE1 null sperm showed a decrease in ATP levels[36]. The loss of ANKRD5 does not change ATP levels in the sperm (Fig. 6E). This result indicates that ANKRD5 can interact with N-DRC components without impacting energy metabolism. The mitochondrial sheath is located in the midpiece of sperm, where mitochondria play a key role in energy production and serve as a significant source of reactive oxygen species (ROS)[42, 43]. A balanced level of ROS is crucial for several essential sperm functions, including motility, capacitation, the acrosome reaction, fertilization, and hyperactivation[44, 45, 46]. Mitochondrial membrane potential (MMP) is correlated with sperm motility[47]. Decreased MMP (depolarization) indicates mitochondrial dysfunction, while increased MMP (hyperpolarization) leads to excess ROS. MMP in spermatozoa was assessed using TMRM staining and fluorescence imaging (Fig. S5 A and B). ROS levels were measured using DCFH-DA staining, followed by fluorescence imaging (Fig. S5 C and D). We found no significant differences between Ankrd5 null sperm and control. Thus, ANKRD5 deficiency did not significantly affect sperm ATP levels and mitochondrial function.

Added in the revised version: Absence of ANKRD5 does not affect energy metabolism.

Earlier version: Absence of ANKRD5 does not affect energy metabolism. (A) The differentially expressed proteins of Ankrd5+/– and Ankrd5+/- were identified by 4D-SmartDIA. (B) Heatmap of relative protein abundance changes between control and knockout mouse sperm. (C) Differences in the expression of N-DRC protein components identified by mass spectra. *P < 0.05, Student’s t test; error bars represent SE (n = 3). (D) GSEA analysis of glycolysis and gluconeogenesis. (E) Measured levels of ATP between wild-type and Ankrd5 null sperm. Student’s t test; error bars represent SE (n = 3).

Revised version: (A) The differentially expressed proteins of Ankrd5+/– and Ankrd5+/- were identified by 4D-SmartDIA. (B) Heatmap of relative protein abundance changes between control and knockout mouse sperm. (C) Differences in the expression of N-DRC protein components identified by mass spectra. *P < 0.05, Student’s t test; error bars represent SE (n = 3). (D) GSEA analysis of glycolysis and gluconeogenesis. (E) Measured levels of ATP between wild-type and Ankrd5 null sperm. Student’s t test; error bars represent SE (n = 3).

Added in the revised version: The ANKRD5-KO sperm axoneme retains ”9+2” configuration

Removed in the revised version: The knockout of Ankrd5 does not affect the overall structure of DMT

Earlier version: The absence of certain DRC components has been shown to alter the expression of other N-DRC proteins. For instance, DRC2/3/4 expression is reduced in DRC1 mutant mice (27), while DRC1/3/5/11 expression is downregulated in Chlamydomonas DRC2 mutants (12). To assess the impact of ANKRD5 deletion on sperm protein composition, we conducted quantitative proteomics using 4D-SmartDIA. Mass spectrometry identified 4,880 quantifiable proteins, of which 126 were differentially expressed (1.5 fold change), including 10 upregulated and 116 downregulated proteins (Fig. 6 A and B). DRC components were not included in the differentially expressed proteins list. In order to better understand the expression changes of these key proteins, the intensity data of DRC components were extracted and normalized, then performed the T-test. Due to technical limitations, DRC6 and DRC12 were not detected. However, the analysis revealed a significant difference in DRC11(IQCA1). As no effective commercial antibody was available, DRC11 was not tested using Western blottign (Fig.6C). TCTE1 deficiency has been associated with abnormalities in glycolysis, leading to decreased ATP levels and reduced sperm motility (29). We conducted a Gene Set Enrichment Analysis (GSEA) focusing on this term, but the results were not significant, which is consistent with our ATP measurement results (Fig.6D and E).

Revised version: The absence of certain DRC components has been shown to alter the expression of other N-DRC proteins. For instance, DRC2/3/4 expression is reduced in DRC1 mutant mice[26], while DRC1/3/5/11 expression is downregulated in Chlamydomonas DRC2 mutants[12]. To assess the impact of ANKRD5 deletion on sperm protein composition, we conducted quantitative proteomics using 4D-SmartDIA. Mass spectrometry identified 4,880 quantifiable proteins, of which 126 were differentially expressed (1.5 fold change), including 10 upregulated and 116 downregulated proteins (Fig. 6 A and B). DRC components were not included in the differentially expressed proteins list. In order to better understand the expression changes of these key proteins, the intensity data of DRC components were extracted and normalized, then performed the T-test. Due to technical limitations, DRC6 and DRC12 were not detected. However, the analysis revealed a significant difference in DRC11(IQCA1). As no effective commercial antibody was available, DRC11 was not tested using Western blotting (Fig.6C). TCTE1 deficiency has been associated with abnormalities in glycolysis, leading to decreased ATP levels and reduced sperm motility[36]. We conducted a Gene Set Enrichment Analysis (GSEA) focusing on this term, but the results were not significant, which is consistent with our ATP measurement results (Fig.6D and E). To further investigate the RS, IDA, and ODA structures of the axonemes, we conducted immunofluorescence assays in both Ankrd5-/- mice and the control group. No significant differences were detected between the two groups (Fig. S6).

Earlier version: The sperm axoneme follows a 9+2 structural arrangement, composed of nine outer microtubule doublets surrounding two central microtubules. Structural defects in this 9+2 arrangement may impair motility, but transmission electron microscopy of the sperm flagellum did not reveal any structural defects in Ankrd5-/- mice (Fig. 4E). To further explore the effects of Ankrd5 knockout on each component of mouse sperm axoneme, we collected about 160 tilt series of Ankrd5 null sperm axoneme using cryo-focused ion beam (cryo-FIB) and cryo-electron tomography (cryo-ET). Data pre-processing and reconstruction were performed by Warp and AreTOMO (43–45) (Fig. S5A). In the original tomograms obtained, we found that the overall “9+2” structure of sperm axonemes remained intact regardless of side view or top view, and there was no significant difference from that of WT mouse sperm axonemes (Fig. 7A) (46), which was similar to TEM results. Among the 89 tomograms with good quality, we carried out manual particle selection of DMT fiber and applied RELION, M, and other software to carry out sub-tomogram average analysis (STA) (47–49). The 8nm-repeats DMT structure of Ankrd5 null sperm axoneme with a resolution of 24 Å was obtained (Fig. S5B and Fig. S6), and the overall structure of its A and B tubes was complete (Fig. 7 B and D). We compared the DMT density map of Ankrd5-/- mouse sperm with that of WT mouse sperm (EMD-35210/35211) and found that there was no significant difference between them, except for some difference in density between A05 of tube A and B09 and B10 of tube B (Fig. 7C and D) (46). This may be due to the difference between the density maps of 8nm-repeats DMT and 16nm-repeats DMT itself. The known 16nm-repeats mouse sperm DMT model was fitted into two density maps, and all 16nm-repeats MIPs including tubulins could be well fitted. We must emphasize that although there appears to be no significant difference on overall DMT structure between the wild-type and Ankrd5 null sperm, we observed notable sample heterogeneity when manually selecting particles for Ankrd5 null sperm. This resulted in a significantly reduced quantity of particles with the 96 nm periodicity, which is distinctly different from the WT results. We were able to detect the N-DRC structure in WT sperm, but we failed to find the density of N-DRC adjacent to RS3 in Ankrd5 null sperm (50) (Fig. S6).

Revised version: The sperm axoneme follows a typical 9+2 structural arrangement, composed of nine outer DMTs surrounding two central microtubules. Structural defects in this 9+2 arrangement may impair motility, but transmission electron microscopy of the sperm flagellum did not reveal any structural defects in Ankrd5-/- mice (Fig. 4E). To further explore the effects of Ankrd5 knockout on each component of mouse sperm axoneme, we collected about 160 tilt series of Ankrd5 null sperm axoneme using cryo-focused ion beam (cryo-FIB) and cryo-electron tomography (cryo-ET) (Fig. S7). Data pre-processing and reconstruction were performed by Warp and AreTOMO[48, 49, 50] (Fig. S8). In the original tomograms obtained, we found that the overall “9+2” structure of sperm axonemes remained intact regardless of side view or top view, and there was no significant difference from that of WT mouse sperm axonemes (Fig. 7A)[51], which was similar to TEM results. In addition, the presence of DMT accessory structures RS and ODA was observed in both WT and ANKRD5-KO tomograms (Fig. 7A and Fig. S9). Among the 89 tomograms with good quality, we carried out manual particle selection of DMT fiber and applied RELION, Dynamo and Warp/M to carry out sub-tomogram averaging (STA) analysis[52, 53, 54]. The 8nm-repeats DMT structure of ANKRD5-KO sperm axoneme with a resolution of 24 Å was obtained (Fig. S7B and Fig. S8), and the overall structure of its A and B tubes was complete (Fig. 7 B and D). We compared the DMT density map of ANKRD5-KO mouse sperm with that of WT mouse sperm (Modified from EMD-35229) and found that there was no significant difference between them, except for some difference in density near A05 of tube A and B10 of tube B (Fig. 7C and D)[51]. The known mouse sperm DMT model was fitted into two density maps, and all 16nm-repeats MIPs including tubulins could be well fitted. This suggests that the loss of ANKRD5 has no significant effect on the overall structure of axoneme and the component proteins of DMT.

Added in the revised version: The overall structure of Ankrd5-KO mouse sperm DMT.

Earlier version: The overall structure of Ankrd5-KO mouse sperm DMT. (A) Side view and top sectional view of WT/Ankrd5-/- mouse sperm axoneme are shown in the tomogram slices. (B) The cryo-EM map of Ankrd5-/- mouse sperm DMT with an 8 nm repeat was obtained by sub-tomogram analysis. (C) Loss of density in Ankrd5-/- DMT structure. The transverse sectional view of DMT is shown. The lost density (khaki color) was obtained by subtracting the density map of Ankrd5-/- DMT from that of the WT DMT. (D) The model of 16nm-repeats WT DMT (PDB: 8I7O) was fitted in the 16nm repeat WT DMT map and Ankrd5-/- DMT map.

Revised version: (A) Side view and top sectional view of WT/Ankrd5-/- mouse sperm axoneme are shown in the tomogram slices. DMT and ODA are marked with white dashed lines and white arrows, respectively. (B) The cryo-EM map of Ankrd5-/- mouse sperm DMT with an 8 nm repeat was obtained by sub-tomogram analysis. (C) Loss of density in Ankrd5-/- DMT structure. The transverse sectional view of DMT is shown. The lost density (khaki color) was obtained by subtracting the density map of Ankrd5-/- DMT from that of the WT DMT. (D) The model of 16nm-repeats WT DMT (PDB: 8I7O) was fitted in the 8nm repeat WT DMT map and Ankrd5-/- DMT map. The 8nm repeats DMT density map was obtained by summing two 16nm repeats DMTs that were staggered 8nm apart from each other.

Added in the revised version: ANKRD5 depletion may impair buffering effect between adjacent DMTs in the axoneme

Added in the revised version: During particle picking of DMT fibers, we observed that transverse sections of axonemal DMT particles from ANKRD5-KO sperm differ markedly from those in WT sperm. Although both A- and B-tubes were visible in both samples, the DMTs in ANKRD5-KO sperm showed a more irregular profile. In WT sperm, DMTs typically appeared circular, whereas ANKRD5-KO DMTs seemed to be extruded as polygonal. (Fig. S9B,D). Notably, ANKRD5-KO DMTs seemed partially open at the junction between the A- and B-tubes (Fig. S9B,D).

Added in the revised version: During the STA process, many particles were misaligned or deformed in the classification results, revealing various degrees of deformation—particularly affecting the B-tube (Figure 9, Fig. S9E). We could retain only ∼10% of the DMT particles to obtain the final density map for ANKRD5-KO sperm (Fig. S9E), whereas ∼70% were usable in WT dataset as reported previously[51]. The mutant DMT density map also displayed roughness at its periphery, indicating substantial structural heterogeneity (Fig. S9E). Even after discarding a large fraction of deformed particles, the final density map still showed evident artifacts, implying that although the mutant DMT preserves the fundamental features of both tubes, its shape is highly heterogeneous (Fig. S9E). Furthermore, attempts to classify the 96-nm repeats did not yield a clear density for radial spokes (RSs) (Fig. S9F), indicating that ANKRD5 deficiency may affect the stability of other accessory structures, such as RSs[23, 24, 25]. In the raw tomograms, RSs in ANKRD5-KO sperm appeared less regularly arranged than those in WT (Fig. S9A and C).

Added in the revised version: Most recently, following the submission of this work, ANKRD5 was reported to localize at the head of the N-DRC, simultaneously binding DRC11, DRC7, DRC4, and DRC5[41]. This structural insight agrees with our in vitro findings that ANKRD5 interacts with DRC4 and DRC5 (Fig. 8C-F). However, that study used isolated and purified DMT samples, leaving the precise positioning of ANKRD5 between adjacent axonemal DMTs unconfirmed. We therefore fitted the published structure (PDB entry: 9FQR) into the in situ DMT structure of mouse sperm 96-nm repeats (EMD-27444), revealing that ANKRD5 lies a mere ∼3 nm from the adjacent DMT (Fig. 8G). Notably, the N-DRC is often likened to a “car bumper”, buffering two neighboring DMTs during vigorous axonemal motion. Given the extensive DMT deformation observed in our cryo-ET data (Fig. S9E), we propose that ANKRD5 contributes to this buffering function at the N-DRC. The loss of ANKRD5 may weaken the “bumper” effect and consequently increase structural damage to adjacent DMTs under intense conditions, while also compromising the stability of associated DMT accessory structures[19, 41, 55].

Added in the revised version: Localization of ANKRD5 in sperm axoneme.

Added in the revised version: (A) Localization of N-DRC in 96nm repeats DMT of mammalian sperm (EMD-50664). (B) Comparison of AlphaFold predicted mouse ANKRD5 structures with known bovine ANKRD5 structures (PDB: 9FQR). (C) Localization of ANKRD5 in 96nm repeats DMT map (EMD-50664). (D) Structure and composition of N-DRC in 96nm repeats DMT of mammalian sperm (EMD-50664). (E) The relationship between ANKRD5 and its interacting proteins as shown in the electron microscope density map (EMD-50664). (F) The relationship between ANKRD5 and its interacting proteins as shown in 96nm repeats DMT model of mammalian sperm (PDB: 9FQR). (G) Localization of ANKRD5 in the in-situ mouse sprem 96nm repeats DMT map (EMD-27444).

Added in the revised version: The functional model of ANKRD5. ANKRD5, as a component of the N-DRC connecting adjacent DMTs, enhances the elasticity of the N-DRC.

Added in the revised version: During sperm movement, the flagellar beating causes adjacent DMT to move towards each other, and the intact N-DRC structure provides good cushioning. In ANKRD5 -knockout mice, the absence of ANKRD5 weakens the N-DRC’s buffering capacity. This subjects the axoneme, especially during intense movement, to greater mechanical stress, leading to the deformation of B tube and impaired sperm motility.

Discussion

Earlier version: During natural fertilization, males could ejaculate millions of sperm, but most are expelled by contractions in the female reproductive tract. Freshly ejaculate sperm exhibit activate movements, and only a few hundred activated sperm reach the isthmus of the fallopian tube, where they attach to the mucosal epithelium to form a sperm reservoir, remaining dormant until ovulation (51). Then, upon release from the reservoir, sperm transition to a hyperactivated state to approach the egg. Activation sperm exhibit symmetrical flagellar oscillations and swim along near-linear trajectories, whereas hyperactivated sperm display high-amplitude, asymmetrical flagellar oscillations (52). Activation is the prerequisite for hyperactivation. The outer layer of the COC is rich in hyaluronic acid, which gives it highly viscoelastic (53). Acrosome reaction releases hyaluronidase, which softens the cumulus cell layer and facilitate sperm-egg interaction. Sperm motility plays an important role in penetrate the COC (33). It has been reported that vole sperm can penetrate the zona pellucida of mice and hamsters without acrosome reaction, underscoring the importance of mechanical forces in sperm penetration of the COC (54).

Revised version: During natural fertilization, males could ejaculate millions of sperm, but most are expelled by contractions in the female reproductive tract. Freshly ejaculate sperm exhibit activate movements, and only a few hundred activated sperm reach the isthmus of the fallopian tube, where they attach to the mucosal epithelium to form a sperm reservoir, remaining dormant until ovulation[56]. Then, upon release from the reservoir, sperm transition to a hyperactivated state to approach the egg. Activation sperm exhibit symmetrical flagellar oscillations and swim along near-linear trajectories, whereas hyperactivated sperm display high-amplitude, asymmetrical flagellar oscillations[57]. Activation is the prerequisite for hyperactivation. The outer layer of the COC is rich in hyaluronic acid, which gives it highly viscoelastic[58]. Acrosome reaction releases hyaluronidase, which softens the cumulus cell layer and facilitate sperm-egg interaction. Sperm motility plays an important role in penetrate the COC[35]. It has been reported that vole sperm can penetrate the zona pellucida of mice and hamsters without acrosome reaction, underscoring the importance of mechanical forces in sperm penetration of the COC[59].

Earlier version: N-DRC is critical for sperm motility. The auxiliary factors required for the N-DRC structure to function effectively have not been fully elucidated. As an important structural component of the axoneme, the functionality of the N-DRC relies on protein interactions. ANKRD5 is evolutionarily conserved and features a helix-turn-helix repeat structure known as the ANK domain, which is a common motif involved in protein-protein interactions. (55). Given the high expression of ANKRD5 in the testes and its presence in the LC-MS results of TCTE1, we are concerned about the potential synergistic role of ANKRD5 in the function of N-DRC. Previous studies have found that ANKRD5 plays a critical role in cell adhesion and protrusion, contributing to the normal formation of the Xenopus gastrula (56). Furthermore, Ankrd5 has been identified as differentially expressed in normozoospermic and asthenozoospermic males (57).

Revised version: N-DRC is critical for sperm motility. The auxiliary factors required for the N-DRC structure to function effectively have not been fully elucidated. As an important structural component of the axoneme, the functionality of the N-DRC relies on protein interactions. ANKRD5 is evolutionarily conserved and features a helix-turn-helix repeat structure known as the ANK domain, which is a common motif involved in protein-protein interactions.[60]. Given the high expression of ANKRD5 in the testes and its presence in the LC-MS results of TCTE1, we are concerned about the potential synergistic role of ANKRD5 in the function of N-DRC. Previous studies have found that ANKRD5 plays a critical role in cell adhesion and protrusion, contributing to the normal formation of the Xenopus gastrula[61]. Furthermore, Ankrd5 has been identified as differentially expressed in normozoospermic and asthenozoospermic males[62].

Ankrd5-/- male mice exhibit normal spermatogenesis and are able to mate and produce mating plugs, but failed to produce offspring. Histological and morphological of the Ankrd5-/- appeared normal. The IVF results found that Ankrd5 null sperm could penetrate cumulus cells, but they could not fertilize cumulus-intact eggs, even after granulosa cell were removed. Interestingly, zona pellucida free eggs were successfully fertilized, and the resulting embryos developed normally to the blastocyst stage. Our findings indicate that the absence of ANKRD5 results in male infertility due to the inability of Ankrd5 null sperm to penetrate the zona pellucida during fertilization. Acrosome reaction and sperm motility is important for sperm to penetrate the ZP. We used A23187 to induce the acrosome reaction and found that Ankrd5 null sperm exhibited normal acrosome reaction. However, CASA revealed reduced motility in Ankrd5 null sperm, suggesting that their inability to penetrate the zona pellucida is linked to abnormal motility rather than defective acrosome function. Our sperm migration results support this conclusion.

Earlier version: The sperm axoneme, a 9+2 motor apparatus, is critical for motility. Dynein arms on A-tubules hydrolyze ATP to “walk” along adjacent B-tubules, driving microtubule sliding. The N-DRC converts this sliding into flagellar bending (24, 25). Mass spectrometry identified several N-DRC members, including DRC5/TCTE1, DRC3/LRRC48, DRC7/CCDC135, and DRC4/GAS8. ANKRD5 knockout sperm shows a phenotype similar to TCTE1 knockout sperm, suggesting ANKRD5 may have a similar role to TCTE1. Drc3/Lrrk48 knockout mice can grow up to adult and are capable of mating, but they are unable to produce offspring (28). In chlamydomonas, Drc3/Lrrc48 mutations cause flagellar motility defects (11). Drc7/Ccdc135 knockout mice exhibit short-tailed sperm, leading to male infertility (30), while DRC4/GAS8 (58) is associated with PCD in humans. Immunoprecipitation validated the mass spectrometry results, and GSEA analysis along with ATP measurements confirmed that ANKRD5 deficiency does not impact energy metabolism. Thus, ANKRD5 likely acts as a structural component aiding N-DRC function without involvement in metabolism.

Revised version: The sperm axoneme, a 9+2 motor apparatus, is critical for motility. Dynein arms on A-tubules hydrolyze ATP to “walk” along adjacent B-tubules, driving microtubule sliding. The N-DRC helps convert this sliding into flagellar bending[23, 24]. Mass spectrometry identified several N-DRC members, including DRC5/TCTE1, DRC3/LRRC48, DRC7/CCDC135, and DRC4/GAS8. ANKRD5 knockout sperm shows a phenotype similar to TCTE1 knockout sperm, suggesting ANKRD5 may have a similar role to TCTE1. Drc3/Lrrc48 knockout mice can grow up to adult and are capable of mating, but they are unable to produce offspring[63]. In chlamydomonas, Drc3/Lrrc48 mutations cause flagellar motility defects[11]. Drc7/Ccdc135 knockout mice exhibit short-tailed sperm, leading to male infertility[21], while DRC4/GAS8[64] is associated with PCD in humans. Immunoprecipitation validated the mass spectrometry results, and GSEA analysis along with ATP measurements confirmed that ANKRD5 deficiency does not impact energy metabolism. Thus, ANKRD5 likely acts as a structural component aiding N-DRC function without involvement in metabolism.

Added in the revised version: The recently reported 96nm repeats DMT structure of bovine sperm axoneme provides a better understanding of the location of ANKRD5 (ANKEF1) and its interaction with other proteins[41]. However, the current structure of N-DRC is still incomplete, especially the binding structure between N-DRC and adjacent DMT. In the future, high-resolution structures of N-DRC at the cellular level are still needed to help us further explore the mechanism of interaction between N-DRC and adjacent DMT, which will also provide us with new insights into the motility mechanism of mammalian sperm.

Removed in the revised version: TEM analysis showed no significant difference in the axoneme structure between wild-type and Ankrd5 null sperm, which may be due to the imaging technique’s low resolution. We therefore aim to use Cryo-ET to compare their electron density. Tilt series of Ankrd5 null sperm axonemes were collected by cryo-focused ion beam and cryo-electron tomography, and 160 tomograms were obtained by data pre-processing and reconstruction. However, in the process of manual selection of DMT fibers, we found that they were not as smooth as WT particles. This suggests that the knockout of Ankrd5 may affect the structural stability of the axoneme, although no significant difference was found in tomograms. We obtained the 8nm-repeats DMT in situ structure of 24 Å Ankrd5 null sperm axonemes by sub-tomogram averaging. On the whole, although there is no obvious difference from the DMT density map of WT, there is a little deformation of tubulin, which does not indicate that DMT is damaged. When classifying 96nm-repeats DMT, we obtained a density similar to RS3, indicating that axoneme components such as DMT, CPC, and RS were not significantly affected, but we failed to obtain the corresponding density of N-DRC, which may be the average result of axoneme motility disturbance caused by decreased sperm motility. It is not clear if N-DRC was damaged.

Earlier version: Asthenospermia, characterized by reduced sperm motility, is a major cause of clinical infertility, yet its precise pathogenesis remains unclear. Investigating the infertility observed in Ankrd5-/- mice may provide insights into the underlying causes of asthenospermia. While current birth control pills primarily target women, they can have adverse side effects such as dizziness, chest pain, decreased libido, and vision impairment. As a result, the development of a male contraceptive has become an urgent priority. Currently, no oral contraceptive is available for men, but our study offers new perspectives for male contraceptive research. ANKRD5 may represent a promising target for the development of such contraceptives.

Revised version: Asthenospermia, characterized by reduced sperm motility, is a major cause of clinical infertility, yet its precise pathogenesis remains unclear. Investigating the infertility observed in Ankrd5-/- mice may provide insights into the underlying causes of asthenospermia. While current birth control pills primarily target women, they can have adverse side effects such as dizziness, chest pain, decreased libido, and vision impairment. As a result, the development of a male contraceptive has become an urgent priority. Currently, no oral contraceptive is available for men, but our study offers new perspectives for male contraceptive research.

Added in the revised version: The potential for male contraceptive development arises from ANKRD5’s critical structural role mediated through its ANK domain, which facilitates interaction with the N-DRC complex in sperm flagella. Recent structural evidence suggests the protein’s positively charged surface may engage with glutamylated tubulin in adjacent microtubules[41], presenting a druggable interface. Targeted disruption of this interaction through small-molecule inhibitors could transiently impair sperm motility. Sperm function relies more on ANKRD5 than respiratory cilia, so inhibiting ANKRD5 has less impact on the latter. This makes ANKRD5 a promising drug target. This tissue-specific phenotypic uncoupling is not uncommon among axonemal-associated proteins, such as DNAH17 and IQUB[65, 66].

Materials and Methods

Animals

All animal experiments were approved by the Animal Care and Use Committee of the Beijing Institute of Biological Sciences. C57BL/6 wild-type mouse were obtained from the SPF Breeding Facility of the Animal Center of the Beijing Institute of Biological Sciences. They were maintained on a 12-hour light-dark cycle (from 7:00 am to 7:00 pm) and provided with adequate food and drinking water.

Real-time Quantitative PCR

RNA was extracted from specific tissues using Trizol (Invitrogen, 15596026). The tissues were thoroughly homogenized with a drill and left at room temperature for 5 minutes to completely lyse the tissue. Chloroform was added and vigorously mixed for 15 seconds. Then, the mixture was incubated at room temperature for 3 minutes and centrifuged at 12,000× g for 15 minutes at 4℃ (repeated twice). Isopropanol was added to the supernatant, gently inverted several times, and left at room temperature for 10 minutes. After centrifugation at 12,000 × g for 10 minutes at 4 °C, the solution was discarded. The pellet was mixed with 75% ethanol, vortexed, and centrifuged at 7500 × g for 5 minutes at 4 °C. The supernatant was discarded, and the pellet was dried for 3-5 minutes. Water was added to resuspend the RNA, which was then shaken well at 65°C for 3 minutes and immediately placed on ice. RNA was reverse-transcribed into cDNA using a kit (TaKaRa, RR047A) according to the instruction manual. The qPCR system was set up following the recommended protocol of the kit (TaKaRa, DRR420A), and the relative expression of Ankrd5 was calculated using Gapdh as the internal reference.

Generation of Ankrd5-deficient and ANKRD5-FLAG Mouse

Female wild-type mouse aged between 4 to 6 weeks were intraperitoneally injected with PMSG to promote follicle development. Subsequently, they were injected with hCG (10 IU/mouse) at a 47-hour interval. After mating with male mouse, fertilized eggs were collected as recipients. Ankrd5-/- mouse were generated by co-injecting Cas9 mRNA and sgRNA1, sgRNA2 into the fertilized eggs. The sequences of sgRNA1 and sgRNA2 are as follows: sgRNA1: 5 ‘cctgcccactaagcggcactatc 3’, sgRNA2: cctctcatgatagcgtgtgccag. ANKRD5-FLAG mouse were created by injecting Cas9 mRNA, gRNA, and ANKRD5-FLAG plasmid into fertilized eggs. Please refer to Fig. 1.C and Fig.S2A for specific strategies.

Histological Analysis

The testes and epididymis tissues were dissected and placed in Davidson’s Fluid (Formaldehyde: Ethanol: Glacial acetic acid: H2O = 6:3:1:10). After fixation at 4℃ for 2 hours, the testis was sliced using a blade, and the tissue sample was fully immersed in the fixing solution at 4 °C overnight. The tissue was then dehydrated, and 5µm thick sections were prepared. The sections were spread with warm water at 42 °C and baked overnight at 42 °C. Subsequently, the sections were dewaxed and stained with hematoxylin and eosin (H&E) following standard protocols. Images were captured using an Olympus VS120 microscope.

Earlier version: In vivo fertilization experiments

Revised version: In vitro fertilization experiments

Female wild-type mouse aged 4-6 weeks were intraperitoneally injected with PMSG to induce follicular development, followed by an injection of human chorionic gonadotropin hCG (10IU/mouse) after a 47-hour interval. After mating with male mouse, fertilized eggs were collected from the female mouse. The oocytes were treated with hyaluronidase (Sigma-Aldrich, H3757) for 10 minutes to remove cumulus cells or with Tyrode’s salt solution (Sigma-Aldrich, T1788) for 1 minute to remove the zona pellucida. Sperm were added to TYH droplets containing cumulus-free or ZP-free oocytes, with a final sperm concentration of 1×106 sperm/ml. For intact egg testing, the final sperm concentration was 1×104 sperm/ml. After co-incubation at 37 °C with 5% CO2, the number of two-cell stage embryos was determined after 36 hours, and the number of blastocysts was determined after 3.5 days.

Acrosome reaction analyses

Sperm was placed in HTF medium under conditions of 37 °C and 5% CO2 for 30 min. Part of sperm were incubated with A23187 (21186-5MG-F) at a final concentration of 10 μmol/L, and the other sperm were incubated with DMSO, incubated for 50min, and then fixed with 4% PFA at room temperature for 30min and spread on the glass. The slides were stained with Coomassie brilliant blue solution (0.22% Coomassie Blue R-250, 50% methanol, 10% glacial acetic acid, 40% water) for 10 minutes and washed with water to remove floating color. Images were acquired using an Olympus VS120 microscope. Each scan included five regions, with a count of 100 sperm per region.

Sperm motility analyses

Sperm was placed in HTF medium under conditions of 37 °C and 5% CO2 for 1 hour. After discarding the cauda tissue, sperm motility was analyzed using the CASA system (Version 14 CEROS, Hamilton Thorne Research) with a Slide Warmer (#720230, Hamilton Thorne Research). The following parameters were used: acquisition of 30 frames at a frame rate of 60 Hz, Minimum Contrast: 30, Minimum Cell Size: 4 Pixels, Minimum Static Contrast: 15.

Sperm Head–Tail Separation

Collect the mouse sperm use liquid nitrogen freezing-thawed for several times. Subsequently, the sample was centrifuged at 10,000× g for 5 minutes, and the pellet was resuspended in 200 μL of PBS. The resuspended sample was mixed with 1.8 mL of 90% Percoll solution (sigma, P1644) and then subjected to centrifugation at 15,000 × g for 15 minutes. Following centrifugation, the sperm head accumulated at the bottom while the sperm tail remained at the top of the tube. Sperm head and tail were mixed with PBS 5 times their volume, and centrifuged at 10000 × g for 5 min. Subsequently, the sperm head and tail were washed twice with PBS, followed by lysing using a lysate buffer (PH=7.6, 50mM Tris-HCl, 150mM NaCl, 1%TritonX-100, 0.5% sodium deoxycholate, 0.1% SDS, 2mM EDTA), protease inhibitor cocktail (Roche, 04693116001) was added to the lysate proportionally. Add 1/5 volume of 5× loading (PH=6.8,10% SDS, 25% Glycerol, 1M Tris-HCL, 5% β-mercaptoethanol, and 0.25% bromophenol blue dye). The mixture was boiled at 100 °C for 10 minutes, and the supernatant was collected as the sample after centrifugation.

Separation of different fraction of Sperm

Earlier version: The separation process for different fractions of sperm was performed as previously described (34, 35, 59). Briefly, the sperm collected from the cauda were washed twice with PBS and resuspended in 1% Triton X-100 lysis buffer (50 mM NaCl, 20 mM Tris⋅HCl, pH 7.5, protease inhibitor mixture) with a protease inhibitor mixture. The samples were then incubated at 4 °C for 2 hours and subsequently centrifuged at 15,000× g for 10 minutes. The resulting supernatant was collected as the Triton-soluble fraction. The pellet was resuspended in 1% SDS lysis buffer (75 mM NaCl, 24 mM EDTA, pH 6.0) and incubated at room temperature for 1 hour. After centrifugation, the supernatant was collected as the SDS-soluble fraction. The residual pellet was boiled in sample buffer (66 mM TRIS-HCL, 2% SDS, 10% glycerol and 0.005% Bromophenol Blue) for 10min. The supernatant obtained after centrifugation was collected as the SDS-resistant fraction. Before use, a protease inhibitor cocktail (Roche, 04693116001) was added to the 1% Triton X-100 lysis buffer and the 1% SDS lysis buffer in the appropriate proportion.

Revised version: The separation process for different fractions of sperm was performed as previously described[37, 38, 67]. Briefly, the sperm collected from the cauda were washed twice with PBS and resuspended in 1% Triton X-100 lysis buffer (50 mM NaCl, 20 mM Tris⋅HCl, pH 7.5, protease inhibitor mixture) with a protease inhibitor mixture. The samples were then incubated at 4 °C for 2 hours and subsequently centrifuged at 15,000× g for 10 minutes. The resulting supernatant was collected as the Triton-soluble fraction. The pellet was resuspended in 1% SDS lysis buffer (75 mM NaCl, 24 mM EDTA, pH 6.0) and incubated at room temperature for 1 hour. After centrifugation, the supernatant was collected as the SDS-soluble fraction. The residual pellet was boiled in sample buffer (66 mM TRIS-HCL, 2% SDS, 10% glycerol and 0.005% Bromophenol Blue) for 10min. The supernatant obtained after centrifugation was collected as the SDS-resistant fraction. Before use, a protease inhibitor cocktail (Roche, 04693116001) was added to the 1% Triton X-100 lysis buffer and the 1% SDS lysis buffer in the appropriate proportion.

Transmission electronic microscopy

Transmission electron microscopy (TEM) was performed at the Centre for Electron Microscopy of the National Institute of Biological Sciences, Beijing, following standard protocols. Briefly, the epididymis cauda was taken and fixed overnight at 4℃ using a 2.5% glutaraldehyde fixing solution (Sigma-Aldrich, G5882). After washing three times with PBS for 20 minutes each time, the samples were fixed at room temperature for 1 hour using 1% Osmium tetroxide. Subsequently, the samples were washed three times with PBS for 20 minutes each time. The samples were then dehydrated using a series of ascending acetone solutions by the progressive lowering temperature method, followed by epoxy and resin embedding. The embedded blocks were placed in a drying oven at 45℃ for 12 hours and then at 60℃ for three days. Ultra-thin slices were obtained and stained with 3% uranyl acetate in 70% methanol/H2O, followed by Sato’s lead for 2 minutes. Images were captured using a TECNAI spirit G2 (FEI) transmission electron microscope at 120 Kv.

Western blot

Samples from various mouse tissues were collected and digested with lysate (PH=7.6, 50mM Tris-HCl, 150mM NaCl, 1%TritonX-100, 0.5% sodium deoxycholate, 0.1% SDS, 2mM EDTA). protease inhibitor cocktail (Roche, 04693116001) was added into the lysate proportionally before use. The lysate was then mixed with 1/5 volume of 5× loading buffer (10% SDS, 25% glycerol, 1M Tris-HCl, 5% β-mercaptoethanol, and 0.25% bromophenol blue dye, pH 6.8). The mixture was boiled at 100 °C for 10 minutes, followed by centrifugation to obtain the supernatant as the sample. Proteins were separated by SDS/PAGE and transferred to PVDF membranes. The membranes were blocked with 5% skim milk in TBST (pH 7.6, 20 mM Tris, 150 mM NaCl with 0.05% Tween-20) for 1 hour. Incubation with primary antibody overnight at 4 °C. Washing the membrane three times with TBST for 10 minutes each time, the membranes were incubated with the secondary antibody at room temperature for 1 hour. Subsequently, the membranes were washed three times with TBST. ECL reagents (BIO-RAD, 170-5060 or NCM Biotech, P10300B) were added onto membranes, and signals were detected by XBT X-ray film (Carestream, 6535876).

Immunofluorescence

Collect mouse sperm in preheated PBS at 37 °C for 15 minutes. Discard the cauda tissue and centrifuge the remaining solution at 1000x g for 5 minutes. The epididymis tissue was discarded, centrifuged at 1000× g for 5min, fixed at room temperature with 4% paraformaldehyde (DF0135; Beijing Leagene Biotech) for 30min, spread on the slide, and dried at room temperature. Antigen repair was performed with antigen repair solution (10 mM sodium citrate, 0.05% Tween-20, pH 6.0), cooled to room temperature, and then blocked with ADB (3% BSA, 0.05% TritonX-100) at room temperature for 1 hour. The primary antibody (anti-FLAG M2) was incubated overnight and washed with PBST (PBS with 0.1% Tween-20) three times for 5min each time. The secondary antibodies (Alexa Fluor® 546-conjugated donkey anti-mouse IgG and Alexa Fluor® 647-conjugated donkey anti-rabbit IgG) and Hoechst 33342 (Sigma, #b2261) were incubated with the samples at room temperature for 1 hour away from light. After incubation, the samples were washed three times with PBST (PBS with 0.1% Triton X-100) for 5 minutes each. Confocal microscopy was performed using a Nikon SIM microscope to acquire images.

Identification of interacting proteome by LC-MS

Protein bands were carefully excised from the polyacrylamide gel using a razor blade, minimizing the inclusion of excess blank gel. The excised gel bands were further trimmed into smaller pieces and transferred into 1.5 mL microcentrifuge tubes for processing. The gel pieces were destained three times using 25 mM ammonium bicarbonate (NH₄HCO₃) in 50% methanol, with each destaining step lasting 10 minutes. Subsequently, the gel fragments were washed three times with a solution of 10% acetic acid in 50% methanol, with each wash lasting 1 hour. Afterward, the gel pieces were rinsed twice with water for 20 minutes each time to ensure removal of residual reagents. The gel pieces were then transferred into 0.5 mL microcentrifuge tubes. To dehydrate the gel fragments, 0.5 mL of 100% acetonitrile (ACN) was added, and the tubes were gently inverted several times until the gel pieces turned opaque white. The ACN was subsequently removed, and the gel pieces were dried using a SpeedVac system for 20–30 minutes under mild heating conditions. For protein digestion, the gel pieces were rehydrated with 5–20 μL of trypsin solution (10 ng/Μl) in 50 mM NH₄HCO₃, pH 8.0). The samples were then incubated at 37 °C overnight to allow for proteolysis. Following digestion, the gel pieces were soaked in 25–50 μL of 50% ACN with 5% formic acid (FA) for 30–60 minutes with gentle agitation, taking care not to vortex the samples. The resulting supernatant was transferred to a fresh 0.5 mL microcentrifuge tube. The gel pieces were then subjected to a second extraction using 25–50 μL of 75% ACN with 0.1% FA, again with gentle agitation for 30–60 minutes. The supernatants from both extractions were combined and dried completely using the SpeedVac system under slight heating conditions, ensuring that the peptides were fully recovered for further analysis.

Protein interactions were verified in 293T cells by co-immunoprecipitation

The protein-coding gene was obtained through the NCBI database and subsequently cloned into the pcDNA3.1 vector. The primers used for vector construction are provided in the supplementary materials (Supplementary Material 1). The plasmid was transfected into 293T cells using the jetOPTIMUS® transfection reagent (catalog number: 101000006), and the cells were cultured at 37 °C with 5% CO₂ for 36 hours. After incubation, the culture medium was discarded, and the cells were gently washed once with PBS to remove any residual medium. Cell lysis was performed using RIPA lysis buffer (pH 7.6: 50 mM Tris-HCl, 150 mM NaCl, 1% Triton X-100 or 1% NP-40, 2 mM EDTA). A portion of the lysate was reserved as the input control, while the remaining lysate was divided into two equal parts. One part was used for FLAG immunoprecipitation (IP) and the other for MYC IP. Both were incubated overnight at 4 °C to allow sufficient binding. The beads were washed with RIPA lysis buffer a total of five times, with each wash lasting 3 minutes. To elute the bound proteins, 1X SDS loading buffer was added to the beads, and the samples were denatured by heating at 95 °C for 5 minutes.

Proteomic analysis of whole sperm

The samples were pulverized into a fine powder using liquid nitrogen and then transferred to a tube. A lysis buffer (8 M urea, 1% protease inhibitor cocktail) was added at a 1:4 ratio, and the mixture was sonicated for 3 minutes on ice using a Scientz high-intensity ultrasonic processor. The lysate was centrifuged at 12,000 g for 10 minutes at 4 °C, and the supernatant was collected. Protein concentration was measured using a BCA assay kit. For digestion, proteins were treated with 5 mM DTT at 56 °C for 30 minutes to reduce disulfide bonds, then alkylated with 11 mM iodoacetamide for 15 minutes in the dark. The sample was diluted with 200 mM TEAB to lower the urea concentration to under 2 M. Trypsin was added at ratio of 1:50 for an overnight digestion, followed by a second digestion at a 1:100 ratio for 4 hours. Peptides were purified using a Strata X SPE column. For mass spectrometry analysis, peptides were dissolved in solvent A (0.1% formic acid, 2% acetonitrile in water) and separated on a 25 cm home-made reversed-phase column with a 100 μm diameter. A gradient elution from 6% to 80% solvent B (0.1% formic acid, 90% acetonitrile) over 20 minutes was used at a flow rate of 700 nL/min with an EASY-nLC 1200 UPLC system. Peptide separation was followed by analysis on an Orbitrap Exploris 480 mass spectrometer, with full MS scans at a 60,000 resolution and MS/MS scans at 15,000 resolution using HCD at 27% NCE. DIA data were processed using the DIA-NN software. Spectra were searched against the Mus_musculus_10090_SP_20231220.fasta database, with Trypsin/P specified as the enzyme and up to one missed cleavage allowed. Fixed modifications included N-terminal methionine removal and carbamidomethylation of cysteine residues. The false discovery rate (FDR) was controlled below 1%.

Measurement of Sperm ATP level

Collect mouse sperm in preheated PBS at 37 °C for 15 minutes. The cauda tissue was discarded, and the remaining solution was centrifuged at 1000× g for 5 minutes. The pellets were then washed twice with PBS. After sperm count, 1×107 sperm per well were added to 96-well plates, followed by adding 30μL of lysate (Promega, G7570). The level of ATP was measured using a MicroPlate Spectrophotometer (TECAN) after shaking the plates on a shaker for 20 minutes away from light.

Evaluation of mitochondrial activity of spermatozoa using TMRM

Earlier version: Sperm were collected from the epididymis cauda and mitochondrial membrane potential was detected by TMRM (Invitrogen, I34361) (60). Briefly, collect mouse sperm in preheated PBS at 37 °C for 15 minutes. Discard the cauda tissue and centrifuge the remaining solution at 300x g for 5 minutes. Wash the pellets twice with DPBS. Incubate the washed sperm with TMRM for 30 minutes at 37℃ and 5% CO2. After incubation, wash the sperm twice with DPBS to remove any unbound dye. Resuspend the pellet and mix it with DAPI, then spread on the slide. Capture fluorescence images using the DAPI filter (405nm) and RFP filter (560nm).

Revised version: Sperm were collected from the epididymis cauda and mitochondrial membrane potential was detected by TMRM (Invitrogen, I34361)[68]. Briefly, collect mouse sperm in preheated PBS at 37 °C for 15 minutes. Discard the cauda tissue and centrifuge the remaining solution at 300x g for 5 minutes. Wash the pellets twice with DPBS. Incubate the washed sperm with TMRM for 30 minutes at 37℃ and 5% CO2. After incubation, wash the sperm twice with DPBS to remove any unbound dye. Resuspend the pellet and mix it with DAPI, then spread on the slide. Capture fluorescence images using the DAPI filter (405nm) and RFP filter (560nm).

Evaluation of ROS level

Spermatozoa were obtained from the cauda epididymis and incubated for 50 minutes in preheated TYH medium (manufacturer, catalog number) with or without Rosup. Subsequently, the samples underwent a gentle centrifugation at 300Xg for 4 minutes, followed by resuspension in 1 mL of DCFDA mix. The DCFDA mix was prepared by diluting DCFH-DA (Beyotime, S0033S) in a 1:1000 ratio with TYH medium. After an incubation at 37 °C with 5% CO2 for 20 minutes, inversion mixing was performed every 5 minutes to ensure proper distribution. The samples were then subjected to another centrifugation step at 300Xg for 4 minutes to remove the supernatant. Pellets were washed three times with DPBS to eliminate any residual DCFH-DA that had not penetrated the intracellular compartments. Following this, the pellets were resuspended, mixed with DAPI, and subsequently spread onto slides. Fluorescence images were captured using a DAPI filter (405nm) and an FITC filter (510-580nm).

Sample preparation of mouse sperm axoneme

Earlier version: Freshly extracted sperm were centrifuged at 400 G (Thermo Scientific Legend Micro 17 R) for 5 min at 4 °C. The precipitate per 100μL semen was carefully re-suspended in 100μL of precooled PBS and diluted 5.5-fold with PBS before use. The cryo-EM carrier network (Quantifoil R3.5/1, Au 200 mesh) was discharged for 60 s using Gatan Solarus. Sperm samples were quickly frozen by vitrification using Leica EMGP. Samples diluted in 3μL PBS were immediately water absorbed at 100% relative humidity and 4 °C for 3-5s, and then the frozen samples were placed in a mixture of ethane and methane cooled to -195 °C and stored in liquid nitrogen for cryo-FIB thinning. The cryo-FIB thinning strategy was referred to in the previous work (46).

Revised version: Freshly extracted sperm were centrifuged at 400 G (Thermo Scientific Legend Micro 17 R) for 5 min at 4 °C. The precipitate per 100μL semen was carefully re-suspended in 100μL of precooled PBS and diluted 5.5-fold with PBS before use. The cryo-EM carrier network (Quantifoil R3.5/1, Au 200 mesh) was discharged for 60 s using Gatan Solarus. Sperm samples were quickly frozen by vitrification using Leica EMGP. Samples diluted in 3μL PBS were immediately water absorbed at 100% relative humidity and 4 °C for 3-5s, and then the frozen samples were placed in a mixture of ethane and methane cooled to -195 °C and stored in liquid nitrogen for cryo-FIB thinning. The cryo-FIB thinning strategy was referred to in the previous work[51].

Cryo-ET tilt series acquisition

Earlier version: The grid after the reduction of cryo-FIB was mounted onto Autoloader in Titan Krios G3 (Thermofisher Scientific) 300 KV TEM, equipped with a Gatan K2 direct electron detector (DED) and a BioQuantum energy filter. Tilt series were collected at a magnification of ×42,000, resulting in a physical pixel size of 3.4 Å in K2 DED and 3.4 Å in counting mode. Before data collection, the pre-tilt of the sample was determined visually, and the pre-tilt was set to 10° or -9° to match the pre-determined geometry induced by loading grids. The total dose for each tilt was set to 3.5 electrons per angstrom squared, divided into 10 frames over a 1.2-s exposure, and the tilt range was set to -9° pre-tilt of -66° to +51°, or +10° pre-tilt of -50° to +67°, in steps of 3°, resulting in 40 tilts and 140 electrons per tilt series. The slit width was set to 20 eV, zero loss peaks were refined after the collection of each tilt series, and nominal defocus was set from -1.8 to -2.5 μm. All tilt series used in this study were collected using a beam-image-shift facilitated acquisition scheme based on a dose symmetry strategy using a script developed in-house in SerialEM software (61–63).

Revised version: The grid after the reduction of cryo-FIB was mounted onto Autoloader in Titan Krios G3 (Thermofisher Scientific) 300 KV TEM, equipped with a Gatan K2 direct electron detector (DED) and a BioQuantum energy filter. Tilt series were collected at a magnification of ×42,000, resulting in a physical pixel size of 3.4 Å in K2 DED and 3.4 Å in counting mode. Before data collection, the pre-tilt of the sample was determined visually, and the pre-tilt was set to 10° or -9° to match the pre-determined geometry induced by loading grids. The total dose for each tilt was set to 3.5 electrons per angstrom squared, divided into 10 frames over a 1.2-s exposure, and the tilt range was set to -9° pre-tilt of -66° to +51°, or +10° pre-tilt of -50° to +67°, in steps of 3°, resulting in 40 tilts and 140 electrons per tilt series. The slit width was set to 20 eV, zero loss peaks were refined after the collection of each tilt series, and nominal defocus was set from -1.8 to -2.5 μm. All tilt series used in this study were collected using a beam-image-shift facilitated acquisition scheme based on a dose symmetry strategy using a script developed in-house in SerialEM software[69, 70, 71].

Mouse sperm axoneme in situ data processing

Earlier version: After data collection, all fractioned movies were imported into Warp for essential processing, including motion correction, 2-x Fourier segmentation of super-resolution frames, CTF estimation, masking platinum islands or other high-contrast features, and tilt series generation. Subsequently, the tilt series was automatically aligned using AreTOMO (45, 64). Aligned tilt series were visually inspected in IMOD and any low-quality frames (such as those blocked by the sample stage or grid bars, containing obvious crystalline ice or showing obvious jumps) were removed to create new sets of tilt series in Warp. The new tilt series underwent a second round of AreTOMO alignment. Then, low-quality frames were removed again using the same criteria as in the first round. The new tilt series was then subjected to a third round of automatic alignment, continuing the process until there were no frames to remove. After tilt series alignment, those tilt series that were less than 30 frames or failed were not further processed (43, 45). The alignment parameters of all remaining tilt series were passed back to Warp, and initial tomogram reconstruction was performed in Warp at 27.2 Å pixel size, resulting in a total of 160 Bin8 tomograms (44).

Revised version: After data collection, all fractioned movies were imported into Warp for essential processing, including motion correction, 2-x Fourier segmentation of super-resolution frames, CTF estimation, masking platinum islands or other high-contrast features, and tilt series generation. Subsequently, the tilt series was automatically aligned using AreTOMO[50, 72]. Aligned tilt series were visually inspected in IMOD and any low-quality frames (such as those blocked by the sample stage or grid bars, containing obvious crystalline ice or showing obvious jumps) were removed to create new sets of tilt series in Warp. The new tilt series underwent a second round of AreTOMO alignment. Then, low-quality frames were removed again using the same criteria as in the first round. The new tilt series was then subjected to a third round of automatic alignment, continuing the process until there were no frames to remove. After tilt series alignment, those tilt series that were less than 30 frames or failed were not further processed[48, 50]. The alignment parameters of all remaining tilt series were passed back to Warp, and initial tomogram reconstruction was performed in Warp at 27.2 Å pixel size, resulting in a total of 160 Bin8 tomograms[49].

Earlier version: Among the total 160 tomograms, we used the filament picking tool in Dynamo to manually pick DMT particles from 89 of all tomograms. By selecting the starting and ending points of each DMT fiber, and separating each cutting point by 8 nm along the fiber axis, we were able to obtain 89 sets of DMT particles (65). The 3D coordinates and two of the three Euler angles (except for in-plane rotation) are automatically generated by Dynamo and then transferred back to Warp for exporting sub-tomograms (44, 65, 66).

Revised version: Among the total 160 tomograms, we used the filament picking tool in Dynamo to manually pick DMT particles from 89 of all tomograms. By selecting the starting and ending points of each DMT fiber, and separating each cutting point by 8 nm along the fiber axis, we were able to obtain 89 sets of DMT particles[73]. The 3D coordinates and two of the three Euler angles (except for in-plane rotation) are automatically generated by Dynamo and then transferred back to Warp for exporting sub-tomograms[49, 73, 74].

Earlier version: In RELION 3.1, sub-tomograms were refined, using the ABTT package to transform the RELION star file and Dynamo table file and use Dynamo and/or RELION to jointly generate a mask (47, 49). First, all particles were reconstructed into the box size of 483 voxels with a pixel size of 27.2 Å, and all extracted particles were directly averaged and low-pass filtered at 80 Å to generate an initial reference. Then 3D classification with K = 1 was performed under the constraints of the first two Euler angles (—sigma_tilt 3 and —sigma_psi 3 in RELION), and 3D automatic refinement was performed after 25 iterations. After alignment, manually cleaned the particles in Chimerax-1.6 (46, 67), and then transferred these aligned parameters back to Warp to export the sub-tomograms with a box size of 843 voxels and a pixel size of 13.6 Å. The particles were automatically refined in RELION. After removing any duplicate particles in Dynamo, we transferred these aligned parameters back to Warp to export the sub-tomograms with the box size of 1283 voxels and the pixel size of 6.8 Å. Then it was automatically refined in RELION and M with a final resolution of 24 Å (47–49).

Revised version: In RELION 3.1, sub-tomograms were refined, using the ABTT package to transform the RELION star file and Dynamo table file and use Dynamo and/or RELION to jointly generate a mask[52, 54]. First, all particles were reconstructed into the box size of 483 voxels with a pixel size of 27.2 Å, and all extracted particles were directly averaged and low-pass filtered at 80 Å to generate an initial reference. Then 3D classification with K = 1 was performed under the constraints of the first two Euler angles (—sigma_tilt 3 and —sigma_psi 3 in RELION), and 3D automatic refinement was performed after 25 iterations. After alignment, manually cleaned the particles in Chimerax-1.6[51, 75], and then transferred these aligned parameters back to Warp to export the sub-tomograms with a box size of 843 voxels and a pixel size of 13.6 Å. The particles were automatically refined in RELION. After removing any duplicate particles in Dynamo, we transferred these aligned parameters back to Warp to export the sub-tomograms with the box size of 1283 voxels and the pixel size of 6.8 Å. Then it was automatically refined in RELION and M with a final resolution of 24 Å[52, 53, 54].

GSEA analysis

Earlier version: The protein mass spectrometry utilized three controls and three ANKRD5 KO mice. Dropout data were imputed using KNN, followed by log2 normalization and sorting to obtain the final data list. Mouse gene symbols were converted to Entrez IDs using the org.Mm.eg.db annotation package, and GSEA analysis was performed with the clusterProfiler package (68). The term ‘Glycolysis and Gluconeogenesis’ is obtained from KEGG. Mouse gene sets were retrieved from the MSigDB (Molecular Signature Database) via the msigdbr package, focusing on category C2, which includes various biological processes and pathways. The analysis was conducted in R version 4.4.1 (2024-06-14 ucrt).

Revised version: The protein mass spectrometry utilized three controls and three ANKRD5 KO mice. Dropout data were imputed using KNN, followed by log2 normalization and sorting to obtain the final data list. Mouse gene symbols were converted to Entrez IDs using the org.Mm.eg.db annotation package, and GSEA analysis was performed with the clusterProfiler package[76]. The term ‘Glycolysis and Gluconeogenesis’ is obtained from KEGG. Mouse gene sets were retrieved from the MSigDB (Molecular Signature Database) via the msigdbr package, focusing on category C2, which includes various biological processes and pathways. The analysis was conducted in R version 4.4.1 (2024-06-14 ucrt).

Statistical analysis

All data are presented as the mean ± SEM (n≥3). Statistical analysis was performed using Student’s t-test or one-way ANOVA. GraphPad Prism version 9.4.1 was used for the analysis, and results were considered significant when P < 0.05. Adobe Illustrator 2021 was used for image layout.